The TOR signaling pathway regulates vegetative development, aflatoxin biosynthesis, and pathogenicity in Aspergillus flavus

Author Response

The following is the authors’ response to the original reviews.

Response to Reviewer 1 Comments (PublicReview)

Point 1: First, the authors should provide more convincing data showing that tor and tapA genes are indeed duplicated genes in A. flavus. The authors appeared to use the A. flavus PTS strain as a parental strain for constructing the tor and tapA mutants. If so, the A. flavus CA14 strain (Hua et al., 2007) should be the parental wild-type strain for the A. flavus PTS strain. I did a BLAST search in NCBI for the torA (AFLA_044350) and tapA (AFLA_092770) genes using the most recent CA14 genome assembly sequence (GCA_014784225.2) and only found one allele for each gene: torA on chromosome 7 and tapA on chromosome 3. I could not find any other parts with similar sequences. Even in another popular A. flavus wild-type strain, NRRL3357, both torA and tapA exist as a single allele. Based on the published genome assembly data for A. flavus, there is no evidence to support the idea that tor and tapA exist as copies of each other. Therefore, the authors could perform a Southern blot analysis to further verify their claim. If torA and tapA indeed exist as duplicate copies in different chromosomal locations, Southern blot data could provide supporting results.

Response 1: We thank the reviewer for their insightful observation. Based on the southern blot analysis results presented in Figure 1, we have determined that torA and tapA are single-copy genes. Additionally, we conducted protoplast transformation experiments repeated several times. which revealed that both torA and tapA transformants exhibited ectopic mutations. It is plausible that the deletion of torA and tapA genes may lead to the demise of A. flavus, this phenomenon is consistent with previous studies conducted on the fungus Fusarium graminearum[1].To ensure the rigor of the study, we have retracted the previously incorrect conclusion. We once again express our heartfelt appreciation to the experts for their valuable suggestions.

Fig.

1 Southern blot hybridization analyses of WT, torA, and tapA transformants. (A) The structure diagram of the torA gene. (B) The structure diagram of the tapA gene. (C) Southern blot hybridization analyses of torA gene. (D) Southern blot hybridization analyses of tapA gene.

Point 2: Second, the authors should consider the possibility of aneuploidy for their constructed mutants. When an essential gene is targeted for deletion, aneuploidy often occurs even in a fungal strain without the "ku" mutation, which results in seemingly dual copies of the gene. As the authors appear to use the A. flavus PTS strain having the "ku" mutation, the parental strain has increased genome instability, which may result in enhanced chromosomal rearrangements. So, it will be necessary to Illumina-sequence their tor and tapA mutants to make sure that they are not aneuploidy.

Response 2: Thank you for your comment. Based on the sequencing results of the torA and tapA mutants, it was determined that the torA and tapA genes were still present in both mutants. In this case, it suggests that the torA and tapA genes may have undergone a genetic rearrangement or insertion at a different site in the mutant strains.

Point 3: Furthermore, the genetic nomenclature +/- and -/- should be reserved for heterozygous and homozygous mutants in a diploid strain. As A. flavus is not a diploid strain, this type of description could cause confusion for the readers.

Response 3: Thank you for your suggestion. We acknowledge your concerns about potential confusion caused by using this type of description, and we agree that it is best to avoid any misunderstandings for readers. Therefore, we have decided to remove this part of the content from the manuscript.

Response to Reviewer 2 Comments (PublicReview)

Point 1: However, findings have not been deeply explored and conclusions are mostly are based on parallel phenotypic observations. In addition, there are some concerns for the conclusions.

Response 1: We are grateful for the suggestion. We conduct additional experiments and analyses to provide a more comprehensive understanding and address concerns raised.

Response to Reviewer 3 Comments (PublicReview)

Point 1: The paper by Li et al. describes the role of the TOR pathway in Aspergillus flavus. The authors tested the effect of rapamycin in WT and different deletion strains. This paper is based on a lot of experiments and work but remains rather descriptive and confirms the results obtained in other fungi. It shows that the TOR pathway is involved in conidiation, aflatoxin production, pathogenicity, and hyphal growth. This is inferred from rapamycin treatment and TOR1/2 deletions. Rapamycin treatment also causes lipid accumulation in hyphae. The phenotypes are not surprising as they have been shown already for several fungi. In addition, one caveat is in my opinion that the strains grow very slowly and this could cause many downstream effects. Several kinases and phosphatases are involved in the TOR pathway. They were known from S. cerevisiae or filamentous fungi. The authors characterized them as well with knock-out approaches.

Response 1: Thank you for your comment. The role of the target of rapamycin (TOR) signaling pathway is of fundamental importance in the physiological processes of diverse eukaryotic organisms. Nevertheless, its precise involvement in regulating the developmental and virulent characteristics of opportunistic pathogenic fungi, such as A. flavus, has yet to be fully elucidated. Furthermore, the mechanistic underpinnings of TOR pathway activity specifically in A. flavus remain largely unresolved. Consequently, our study represents a significant contribution as the first comprehensive exploration of the conserved TOR signaling pathway encompassing a majority of its constituent genes in A. flavus.

Response to Reviewer 1 Comments (Recommendations For The Authors)

Point 1: In Table S3, the authors indicated that the Δku70 ΔniaD ΔpyrG::pyrG strain is A. flavus wild-type strain. However, this strain is not a wild-type strain because it seems like a control strain after introducing the pyrG gene into the A. flavus PTS strain (Δku70 ΔniaD ΔpyrG). So please indicate the real wild-type A. flavus strain name to help readers find out its original genome sequence data. Also, the reference for this Δku70 ΔniaD ΔpyrG::pyrG strain is "saved in our lab". This is not an eligible reference. If you use this control strain for the first time in this study, it should be described as "In this study". Otherwise, please indicate the proper reference for which the strain was first used.

Response 1: Thank you for your valuable feedback on our manuscript. We appreciate your attention to detail and the opportunity to clarify the information regarding the strain in Table S3. The A. flavus CA14 strain which produces aflatoxins and large sclerotia was isolated from a pistachio bud in the Wolfskill Grant Experimental Farm (University of Davis, Winters, California, USA)[2]. The A. flavus CA14 strain is the parental wild-type strain for the A. flavus CA14 PTs (Δku70, ΔniaD, ΔpyrG) strain. The recipient strain CA14 PTs has been used satisfactorily in gene knockout and subsequent genetic complementation experiments[3]. In this study, the A. flavus CA14 PTs strain was used as the transformation recipient strain, and the control strain (Δku70, ΔniaD, ΔpyrG::pyrG) created by introducing the pyrG gene into the A. flavus CA14 PTs strain. Refer to previously published literature[4]，this control strain (Δku70, ΔniaD, ΔpyrG::pyrG) was named wild-type strain. Therefore, this control strain was also named wild-type strain in this study. As this control strain is indeed used in this study, we will revise the reference to "In this study" Once again, we appreciate your keen attention to detail and thank you for bringing these issues to our attention.

Response to Reviewer 2 Comments (Recommendations For The Authors)

Point 1: As in response: However, the tor gene in A. flavus exhibited varying copy numbers, as was confirmed by absolute quantification PCR at the genome level (Table S1). However, it is hard to understand Table S1: Estimation of copy number of tor gene in A. flavus toro and sumoo stand for the initial copy number, and the data are figured as the mean {plus minus} 95% confidence limit. CN is copy number. As indicated in the section of Method, using sumo gene as reference, the tor and tapA gene copy number was calculated by standard curve. In Table S1 of WT, for tor gene, CN value is 1412537 compared to 1698243 in tor+/-, for the reference gene sumo,794328 compared to1584893, how these data could support copy gene numbers of tor?

Response 1: Thank you for your suggestion. We understand the confusion with the data presented in Table S1 regarding the copy number estimation of the tor gene in A. flavus. We apologize for not providing a clear explanation for the data in the table. Quantitative real-time PCR (qPCR) is widely used to determine the copy number of a specific gene. It involves amplifying the gene of interest and a reference gene simultaneously using specific primers and probes. By comparing the amplification curves of the gene of interest and the reference gene, you can estimate the relative copy number of the gene.

To address your concern and provide more accurate information, we have re-performed the copy number analysis using southern blot. Southern blot analysis allows for the direct estimation of gene copy number by hybridizing genomic DNA with a specific probe for the gene. This method provides more reliable and accurate results in determining gene copy numbers. The southern blot analysis results are presented in Figure 1.

We appreciate your input and apologize for any confusion caused by the earlier presentation of the data.

Point 2: In response: For the knockout of the FRB domain, we used the homologous recombination method, but because tor genes are double-copy genes, there are also double copies in the FRB domain. Despite our efforts, we encountered challenges in precisely determining the location of the other copy of the tor gene. I could not understand these consistent data, why not for using sequencing?

Response 2: Thank you for your comment. We observed that the torA gene is a single copy. We removed this part of the results to avoid any ambiguity or potential misinterpretation.

Point 3: Response in Due to the large number of genes involved, we did not perform a complementation experiment. If there were no complementation data, how to demonstrate data are solid?

Response 3: Thank you for your important suggestion. We understand that complementation experiments are commonly used to validate gene deletions. Therefore, to ensure the reliability of our data, we have conducted supplementary experiments on specific gene deletions, such as ΔsitA-C and Δppg1-C. Thank you again for your positive comments and valuable suggestions to improve the quality of our manuscript.

References:

(1) Yu F, Gu Q, Yun Y, et al. The TOR signaling pathway regulates vegetative development and virulence in Fusarium graminearum. New Phytol. 2014; 203(1): 219-32.

(2) Hua SS, Tarun AS, Pandey SN, Chang L, Chang PK. Characterization of AFLAV, a Tf1/Sushi retrotransposon from Aspergillus flavus. Mycopathologia. 2007 Feb;163(2):97-104.

(3) Chang PK, Scharfenstein LL, Mack B, Hua SST. Genome sequence of an Aspergillus flavus CA14 strain that is widely used in gene function studies. Microbiol Resour Announc. 2019 Aug 15;8(33):e00837-19.

(4) Zhu Z, Yang M, Yang G, Zhang B, Cao X, Yuan J, Ge F, Wang S. PP2C phosphatases Ptc1 and Ptc2 dephosphorylate PGK1 to regulate autophagy and aflatoxin synthesis in the pathogenic fungus Aspergillus flavus. mBio. 2023 Oct 31;14(5):e0097723.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The target of rapamycin (TOR) signaling pathway is highly conserved and plays a crucial role in diverse biological processes in eukaryotes. However, the underlying mechanism of the TOR pathway in A. flavus remains elusive. In this study, we identified and characterized nine genes encoding various components of the TOR pathway in A. flavus, and investigated their biological function. The FK506-binding protein Fkbp3 and its lysine succinylation are important for aflatoxin production and rapamycin resistance. The Tor kinase plays a pivotal role in the regulation of growth, spore production, aflatoxin biosynthesis, and osmotic and rapamycin stress. As a significant downstream effector molecule of the Tor kinase, the Sch9 kinase regulates aflatoxin B1 (AFB1) synthesis, osmotic and calcium stress response in A. flavus, and this regulation is mediated through its S_TKc, S_TK_X domains, and the ATP binding site at K340. We also showed that the Sch9 kinase may have a regulatory impact on the high-osmolarity glycerol (HOG) signaling pathway. TapA and TipA, the other downstream components of the Tor kinase, play significant roles in regulating sclerotia formation and cell wall stress response in A. flavus. The members of the TapA-phosphatase complexes, SitA and Ppg1, are crucial for various biological processes in A. flavus, including vegetative growth, sclerotia formation, AFB1 biosynthesis, and pathogenicity. Furthermore, we showed that SitA and Ppg1 are involved in regulating lipid droplets (LDs) biogenesis and cell wall integrity (CWI) signaling pathway. In addition, another phosphatase complex, Nem1/Spo7, plays critical roles in hyphal development, conidiation, aflatoxin, and lipid droplets biogenesis. This study provides an important insight into the regulatory network of the TOR signaling pathway and the molecular mechanism of aflatoxin biosynthesis in A. flavus.

Revised version: The target of rapamycin (TOR) signaling pathway is highly conserved and plays a crucial role in diverse biological processes in eukaryotes. However, the underlying mechanism of the TOR pathway in A. flavus remains elusive. In this study, we identified and characterized seven genes encoding various components of the TOR pathway in A. flavus, and investigated their biological function. The FK506-binding protein Fkbp3 and its lysine succinylation are important for aflatoxin production and rapamycin resistance. As a significant downstream effector molecule of the TorA kinase, the Sch9 kinase regulates aflatoxin B1 (AFB1) synthesis, osmotic and calcium stress response in A. flavus, and this regulation is mediated through its S_TKc, S_TK_X domains, and the ATP binding site at K340. We also showed that the Sch9 kinase may have a regulatory impact on the high-osmolarity glycerol (HOG) signaling pathway. TipA, the other downstream component of the TorA kinase, plays a significant role in regulating sclerotia formation and cell wall stress response in A. flavus. The members of the TapA-phosphatase complexes, SitA and Ppg1, are crucial for various biological processes in A. flavus, including vegetative growth, sclerotia formation, AFB1 biosynthesis, and pathogenicity. Furthermore, we showed that SitA and Ppg1 are involved in regulating lipid droplets (LDs) biogenesis and cell wall integrity (CWI) signaling pathways. In addition, another phosphatase complex, Nem1/Spo7, plays critical roles in hyphal development, conidiation, aflatoxin, and lipid droplets biogenesis. This study provides an important insight into the regulatory network of the TOR signaling pathway and the molecular mechanism of aflatoxin biosynthesis in A. flavus.

Introduction

Earlier version: Aspergillus flavus is one of the most important phytopathogenic fungi that frequently contaminates various plants, causing economic losses in agricultural commodities and posing serious health concerns in developing countries [1]. A. flavus is also an opportunistic human pathogenic fungus that can cause aspergillosis in immune-compromised patients [2]. Aflatoxins (AFs), the toxic and carcinogenic secondary metabolites produced by A. flavus, are considered the most potent carcinogen present in nature, causing liver cancer in both humans and animals [3]. The aflatoxin biosynthesis is a complex process involving a series of reactions. Although the majority of the enzyme reactions and genes involved in aflatoxin biosynthesis have been elucidated [4], the related signaling pathway network and specific regulatory mechanisms governing aflatoxin biosynthesis remain elusive. Therefore, exploring the regulatory mechanism of the aflatoxin biosynthesis signaling pathway would provide new insights for preventing A. flavus and aflatoxins contamination.

Revised version: Aspergillus flavus is one of the most important phytopathogenic fungi that frequently contaminate various plants, causing economic losses in agricultural commodities and posing serious health concerns in developing countries [1]. A. flavus is also an opportunistic human pathogenic fungus that can cause aspergillosis in immune-compromised patients [2]. Aflatoxins (AFs), the toxic and carcinogenic secondary metabolites produced by A. flavus, are considered the most potent carcinogen present in nature, causing liver cancer in both humans and animals [3]. The aflatoxin biosynthesis is a complex process involving a series of reactions. Although the majority of the enzyme reactions and genes involved in aflatoxin biosynthesis have been elucidated [4], the related signaling pathway network and specific regulatory mechanisms governing aflatoxin biosynthesis remain elusive. Therefore, exploring the regulatory mechanism of the aflatoxin biosynthesis signaling pathway would provide new insights for preventing A. flavus and aflatoxin contamination.

In eukaryotes, the target of rapamycin (TOR) signaling pathway is highly conserved and plays essential roles in various significant biological processes, including ribosome biosynthesis, cell growth, and autophagy [5]. The serine/threonine-protein kinase Tor serves as a crucial protein component in the TOR signaling pathway, which interacts with other proteins to form multi-protein complexes [6]. In Saccharomyces cerevisiae, TOR exists in two distinct multiprotein functional complexes, namely TORC1 and TORC2 [7]. TORC1-mediated signaling governs cellular growth by modulating various growth-related mechanisms and exhibits sensitivity to rapamycin. Conversely, TORC2-mediated signaling primarily regulates cytoskeletal remodeling but remains unaffected by rapamycin [7]. In yeast, there are two tor genes, tor1 and tor2, whereas in higher eukaryotes such as plants, animals, and filamentous fungi, there is only one tor gene [8–10].

Earlier version: In the fungal kingdom, the TOR signaling pathway has been extensively investigated in three prominent species: the budding yeast S. cerevisiae [11–12], the fission yeast Schizosaccharomyces pombe [13–15], and the human pathogen Candida albicans [16–17]. In S. cerevisiae, macrolide antibiotic rapamycin interacts with the cytosolic protein Fkbp12 to form a complex, and the Fkbp12-rapamycin complex effectively inhibits the activity of the Tor kinase function by binding to the FRB domain [18]. The Tap42 phosphatase complexs serves as the primary target of the Tor kinase. Additionally, the Tap42 protein interacts with the Tip41 protein to collaboratively regulate phosphatases, including PP2A and Sit4 [19–20]. Phosphatase Sit4 can dephosphorylate the transcription factor Gln3, regulating the nitrogen source metabolism pathway [21]. In addition, another immediate effector of TORC1 is the AGC kinase Sch9, which participates in various cellular functions processes, such as ribosome biosynthesis, translation initiation, and entry into G0 phase [22]. The TOR signaling pathway has been documented in various filamentous fungi, including Fusarium graminearum [23–24], Magnaporthe oryzae [25], Phanerochaete chrysosporium [26], Podospora anserine [27], A. nidulans [28], A. fumigatus [29–30], F. fujikuroi [31], Botrytis cinerea [32], F. oxysporum [33], and Mucor circinelloides [34]. However, the TOR signaling pathway has not yet been reported in A. flavus.

Revised version: In the fungal kingdom, the TOR signaling pathway has been extensively investigated in three prominent species: the budding yeast S. cerevisiae [11–12], the fission yeast Schizosaccharomyces pombe [13–15], and the human pathogen Candida albicans [16–17]. In S. cerevisiae, macrolide antibiotic rapamycin interacts with the cytosolic protein Fkbp12 to form a complex, and the Fkbp12-rapamycin complex effectively inhibits the activity of the Tor kinase function by binding to the FRB domain [18]. The Tap42 phosphatase complex serves as the primary target of the Tor kinase. Additionally, the Tap42 protein interacts with the Tip41 protein to collaboratively regulate phosphatases, including PP2A and Sit4 [19–20]. Phosphatase Sit4 can dephosphorylate the transcription factor Gln3, regulating the nitrogen source metabolism pathway [21]. In addition, another immediate effector of TORC1 is the AGC kinase Sch9, which participates in various cellular functions processes, such as ribosome biosynthesis, translation initiation, and entry into the G0 phase [22]. The TOR signaling pathway has been documented in various filamentous fungi, including Fusarium graminearum [23–24], Magnaporthe oryzae [25], Phanerochaete chrysosporium [26], Podospora anserine [27], A. nidulans [28], A. fumigatus [29–30], F. fujikuroi [31], Botrytis cinerea [32], F. oxysporum [33], and Mucor circinelloides [34]. However, the TOR signaling pathway has not yet been reported in A. flavus.

Earlier version: In eukaryotic cells, the TOR signaling pathway and the Mitogen Activated Protein Kinase (MAPK) signaling pathways play an important role in regulating adaptive responses to extra- and intracellular conditions [35]. Several studies have demonstrated that the TOR signaling pathway interacts with various other signaling pathways, such as MAPK and CWI. These pathways engage in crosstalk, forming a complex metabolic network that regulates cellular processes [36–40]. However, the underlying mechanisms of multiple crosstalks between the TOR and other signaling pathways in A. flavus remain unclear. Therefore, we intended to identify the genes associated with the TOR signaling pathway and elucidate their roles and contributions in regulating vegetative development and aflatoxin biosynthesis in A. flavus. Our results demonstrated that the TOR signaling pathway is involved in multiple cellular processes in A. flavus. Our findings revealed a complex interplay between the TOR pathway and other signaling pathways (such as HOG and CWI) in A. flavus, which play a crucial role in regulating cellular growth and survival under environmental stress conditions.

Revised version: In eukaryotic cells, the TOR signaling pathway and the mitogen-activated protein kinase (MAPK) signaling pathways play an important role in regulating adaptive responses to extra- and intracellular conditions [35]. Several studies have demonstrated that the TOR signaling pathway interacts with various other signaling pathways, such as MAPK and cell wall integrity (CWI) pathway. These pathways engage in crosstalk, forming a complex metabolic network that regulates cellular processes [36–40]. However, the underlying mechanisms of multiple crosstalks between the TOR and other signaling pathways in A. flavus remain unclear. Therefore, we intended to identify the genes associated with the TOR signaling pathway and elucidate their roles and contributions in regulating vegetative development and aflatoxin biosynthesis in A. flavus. Our results demonstrated that the TOR signaling pathway is involved in multiple cellular processes in A. flavus. Our findings revealed a complex interplay between the TOR pathway and other signaling pathways (such as HOG and CWI) in A. flavus, which play a crucial role in regulating cellular growth and survival under environmental stress conditions.

Results

Rapamycin exhibits inhibitory effects on the growth, sporulation, sclerotia formation, and aflatoxin production in A. flavus

Earlier version: Rapamycin, a secondary metabolite synthesized by Streptomyces hygroscopicus, exhibits efficacy in suppressing specific filamentous fungi. In F. graminearum, rapamycin has a significant inhibitory impact on growth and asexual reproduction [23]. To investigate the effects of rapamycin on vegetative development, sclerotia formation, and secondary metabolism in A. flavus, we conducted an assay to determine the sensitivity of the wild-type strain to rapamycin. We observed that the wild-type strain exhibited a high sensitivity to rapamycin. We found that 100 ng/mL of rapamycin significantly inhibited mycelial growth and conidia formation (Fig.1A, 1D, 1E). More importantly, the wild-type strain exhibited a significant reduction in sclerotia formation and aflatoxin synthesis (Fig.1B, 1C, 1F, 1G). Additionally, rapamycin resulted in a significant decrease in spore germination rate and sparser conidiophores (Fig. S1A, S1B). These results illustrated that the TOR signaling pathway may be involved in regulating the growth, development, and secondary metabolism in A. flavus. Calcofluor white is a chemifluorescent blue dye commonly employed for its nonspecific binding affinity to fungal chitin. Fluorescence microscopy analysis revealed a notable reduction in the blue fluorescence on the cell wall following treatment with rapamycin (Fig.S1C). The intracellular lipid droplets within the hyphae were labeled with BODIPY. We observed that the hyphae treated with rapamycin contained more lipid droplets compared to the control group (Fig. S1D). These findings suggest that the TOR pathway may play important roles in maintaining cell wall integrity and facilitating the biogenesis of lipid droplets.

Revised version: Rapamycin, a secondary metabolite synthesized by Streptomyces hygroscopicus, exhibits efficacy in suppressing specific filamentous fungi. In F. graminearum, rapamycin has a significant inhibitory impact on growth and asexual reproduction [23]. To investigate the effects of rapamycin on vegetative development, sclerotia formation, and secondary metabolism in A. flavus, we conducted an assay to determine the sensitivity of the wild-type strain to rapamycin. We observed that the wild-type strain exhibited a high sensitivity to rapamycin. We found that 100 ng/mL of rapamycin significantly inhibited mycelial growth and conidia formation (Fig. 1A, 1D, 1E). More importantly, the wild-type strain exhibited a significant reduction in sclerotia formation and aflatoxin synthesis (Fig. 1B, 1C, 1F, 1G). Additionally, rapamycin resulted in a significant decrease in spore germination rate and sparser conidiophores (Fig. S1A, S1B). These results illustrated that the TOR signaling pathway may be involved in regulating the growth, development, and secondary metabolism of A. flavus. Calcofluor white is a chemifluorescent blue dye commonly employed for its nonspecific binding affinity to fungal chitin. Fluorescence microscopy analysis revealed a notable reduction in the blue fluorescence on the cell wall following treatment with rapamycin (Fig. S1C). The intracellular lipid droplets within the hyphae were labeled with BODIPY. We observed that the hyphae treated with rapamycin contained more lipid droplets compared to the control group (Fig. S1D). These findings suggest that the TOR pathway may play important roles in maintaining cell wall integrity and facilitating the biogenesis of lipid droplets.

Impacts of rapamycin on the growth, sporulation, sclerotia formation,and aflatoxin production of Aspergillus flavus.

(A) Colony morphology of the WT strain cultured on PDA medium amended with different concentrations of rapamycin at 37°C for 5 days.

(B) Colony phenotype of the WT strain grown on CM medium amended with different concentrations of rapamycin at 37°C for 7 days.

(C) TLC assay of AFB1 production by the WT strain cultured in YES liquid media amended with different concentrations of rapamycin at 29°C for 6 days.

(D) Statistical analysis of the colony diameter by the WT strain treated with different concentrations of rapamycin described in (A).

(E) Conidial quantification of the WT strain treated with different concentrations of rapamycin as mentioned in (A).

(F) Quantitative analysis of sclerotium formation by the WT strain treated with different concentrations of rapamycin described in (B).

(G) Relative quantification of AFB1 production by the WT strain treated with different concentrations of rapamycin as mentioned in (C)

(H) Relative expression levels of conidia-related genes in WT strain treated with 100 ng/mL rapamycin after 3, 6, 9 h.

(I) Relative expression levels of sclerotia-related genes in WT strain treated with 100 ng/mL rapamycin after 3, 6, 9 h.

(J) Relative expression levels of aflatoxin biosynthesis regulatory and structural genes in WT strain treated with 100 ng/mL rapamycin after 3, 6, 9 h.

To investigate the potential regulatory role of the TOR pathway in modulating the sporulation, sclerotia formation, and aflatoxin biosynthesis in A. flavus, we examined the expression levels of conidia-related, sclerotia formation-related, and aflatoxin-related genes after 3, 6, and 9 hours of treatment with rapamycin. To determine the effects of rapamycin on conidiation, qRT-PCR was conducted to assess the transcript levels of two conidia-related genes, abaA and brlA. We also conducted qRT-PCR to evaluate the expression levels of three regulators (nsdC, nsdD, and sclR) involved in sclerotia development, and genes involved in AFB1 biosynthesis, such as the regulatory gene aflS, as well as several structural genes (aflC and aflQ). The strains treated with rapamycin exhibited a significant reduction in the expression levels of genes related to sporulation, sclerotia formation, and aflatoxin synthesis (Fig. 1H, 1I, 1J). Overall, these results indicated that the TOR pathway plays a crucial role in regulating the growth, sporulation, sclerotia formation, and aflatoxin biosynthesis in A. flavus by modulating the expression of these genes.

The impact of Fkbp3 and its lysine succinylation on AFB1 biosynthesis and rapamycin resistance in A. flavus

Earlier version: FK506-binding proteins (Fkbps) belong to a highly conserved immunophilin family, which are involved in various cellular functions, such as protein folding, immunosuppression, signaling transduction, and transcription [41]. Fkbps function as molecular switches by binding to target proteins and inducing conformational changes [41]. In S. cerevisiae, Fkbp12 and rapamycin interact to form a complex that exerts a negative regulatory effect on the activity of the Tor kinase [42]. The fkbp12 gene homologue fprA in A. nidulans shares high levels of homology with that from S. cerevisiae [28]. According to the protein sequence of FprA from A. nidulans, four genes encoding putative proteins with FK506 binding domain, named fkbp1 (AFLA_126880), fkbp2 (AFLA_087480), fkbp3 (AFLA_010910), and fkbp4 (AFLA_128200), have been identified from the A. flavus genome. We generated four fkbp knockout strains using homologous recombination, and all fkbp disruption strains were confirmed by PCR and sequencing analysis (Fig.S2). Phenotypic analyses showed that the mycelial growth and conidia formation of all the fkbp mutants were comparable to the wild-type strain on the PDA medium (Fig. 2A), suggesting that Fkbps have minimal influence on the process of vegetative development. We evaluated AFB1 synthesis utilizing thin layer chromatography (TLC), and the findings demonstrated a noteworthy reduction in AFB1 synthesis upon deletion of Fkbp3, in comparison to the wild-type strain (Fig.2B, 2D). Additionally, the quantity of sclerotia exhibited a significant decrease in all fkbp mutants compared to the wild-type strain (Fig.S3A, S3B). The aforementioned data indicated that Fkbp3 plays a crucial role in sclerotia formation and aflatoxin biosynthesis in A. flavus.

Revised version: FK506-binding proteins (Fkbps) belong to a highly conserved immunophilin family, which are involved in various cellular functions, such as protein folding, immunosuppression, signaling transduction, and transcription [41]. Fkbps function as molecular switches by binding to target proteins and inducing conformational changes [41]. In S. cerevisiae, Fkbp12 and rapamycin interact to form a complex that exerts a negative regulatory effect on the activity of the Tor kinase [42]. The fkbp12 gene homologue fprA in A. nidulans shares high levels of homology with that from S. cerevisiae [28]. According to the protein sequence of FprA from A. nidulans, four genes encoding putative proteins with FK506 binding domain, named fkbp1 (AFLA_126880), fkbp2 (AFLA_087480), fkbp3 (AFLA_010910), and fkbp4 (AFLA_128200), have been identified from the A. flavus genome. We generated four fkbp knockout strains using homologous recombination, and all fkbp disruption strains were confirmed by PCR and sequencing analysis (Fig. S2). Phenotypic analyses showed that the mycelial growth and conidia formation of all the fkbp mutants were similar to the wild-type strain on the PDA medium (Fig. 2A), suggesting that Fkbps have minimal influence on the process of vegetative development. We evaluated AFB1 synthesis utilizing thin layer chromatography (TLC), and the findings demonstrated a noteworthy reduction in AFB1 synthesis upon deletion of Fkbp3, in comparison to the wild-type strain (Fig. 2B, 2D). Additionally, the quantity of sclerotia exhibited a significant decrease in all fkbp mutants compared to the wild-type strain (Fig. S3A, S3B). The aforementioned data indicated that Fkbp3 plays a crucial role in sclerotia formation and aflatoxin biosynthesis in A. flavus.

Disruption of fkbp3 significantly increases the resistance of A. flavus to rapamycin and FK506.

(A) Phenotype of the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) grown on PDA amended with rapamycin and FK506 at 37°C for 5 days.

(B) TLC analysis of AFB1 production by the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) cultured in YES liquid medium at 29°C for 6 days.

(C) The growth inhibition rate of the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) under rapamycin and FK506 stress.

(D) AFB1 quantitative analysis of the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) as described in (B).

(E) Phenotype of the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

(F) The growth inhibition rate of the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) under rapamycin stress.

(G) TLC assay of AFB1 production by the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) cultured in YES liquid medium at 29°C for 6 days.

(H) Relative quantification of AFB1 production in the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) as mentioned in (G).

Earlier version: To determine the specific Fkbp responsible for rapamycin sensitivity, all mutant strains were cultured on PDA solid plates amended with 100 ng/mL rapamycin. The strain with disrupted fkbp3 showed significant resistance to rapamycin, while the deletion mutants of fkbp1, fkbp2, and fkbp4 did not exhibit any alteration in resistance to rapamycin (Fig.2A, 2C). These results demonstrated that Fkbp3 serves as the principal target for rapamycin in A. flavus. In the FK506-sensitivity analysis, we observed a significant suppression of mycelial growth and conidia formation in A. flavus with the addition of 10 ng/mL FK506 to the PDA medium. Additionally, the deletion of fkbp3 led to the heightened resistance to FK506 compared to the wild-type strain (Fig.2A, 2C). These findings suggested that Fkbp3 is mainly involved in the modulation of FK506 resistance in A. flavus.

Revised version: To determine the specific Fkbp responsible for rapamycin sensitivity, all mutant strains were cultured on PDA solid plates amended with 100 ng/mL rapamycin. The strain with disrupted fkbp3 showed significant resistance to rapamycin, while the deletion mutants of fkbp1, fkbp2, and fkbp4 did not exhibit any alteration in resistance to rapamycin (Fig. 2A, 2C). These results demonstrated that Fkbp3 serves as the principal target for rapamycin in A. flavus. In the FK506-sensitivity analysis, we observed a significant suppression of mycelial growth and conidia formation in wild-type strain with the addition of 10 ng/mL FK506 to the PDA medium. Additionally, the deletion of fkbp3 led to the heightened resistance to FK506 compared to the wild-type strain (Fig. 2A, 2C). These findings suggested that Fkbp3 is mainly involved in the modulation of FK506 resistance in A. flavus.

Earlier version: Previous studies have shown that lysine succinylation plays a significant role in aflatoxin production and pathogenicity in A. flavus [43]. Based on the analysis of our succinylome data, we have identified six reliable succinylation sites (K5, K19, K40, K42, K55, and K65) on the Fkbp3 protein. To identify the function of Fkbp3 succinylation sites, site-directed mutations were constructed according to the homologous recombination strategy. Compared to the wild-type strain, all point mutants exhibited a comparable phenotype in vegetative growth and conidiation. Interestingly, we observed that the K19A mutation resulted in increased resistance to rapamycin (Fig. 2E, 2F). Additionally, the K19A strain exhibited a significant reduction in the production of AFB1 compared to the wild-type strain (Fig. 2G, 2H). These findings suggested that the succinylation of Fkbp3 at K19 plays a crucial role in rapamycin resistance and aflatoxin biosynthesis. Furthermore, we found that various point mutants, including K40A, K42A, and K55A, exhibited reduced levels of AFB1 production in comparison to the wild-type strain (Fig. 2G, 2H), suggesting that the succinylation of Fkbp3 at other sites also contributes to AFB1 biosynthesis in A. flavus. The above results indicated that the K19 site likely plays a more prominent role in Fkbp3 succinylation compared to other sites.

Revised version: Previous studies have shown that lysine succinylation plays a significant role in aflatoxin production and pathogenicity in A. flavus [43]. Based on the analysis of our succinylome data, we have identified six reliable succinylation sites (K5, K19, K40, K42, K55, and K65) on the Fkbp3 protein. To identify the function of Fkbp3 succinylation sites, site-directed mutations were constructed according to the homologous recombination strategy. Compared to the wild-type strain, all point mutants exhibited a similar phenotype in vegetative growth and conidiation. Interestingly, we observed that the K19A mutation resulted in increased resistance to rapamycin (Fig. 2E, 2F). Additionally, the K19A strain exhibited a significant reduction in the production of AFB1 compared to the wild-type strain (Fig. 2G, 2H). These findings suggested that the succinylation of Fkbp3 at K19 plays a crucial role in rapamycin resistance and aflatoxin biosynthesis. Furthermore, we found that various point mutants, including K40A, K42A, and K55A, exhibited reduced levels of AFB1 production in comparison to the wild-type strain (Fig. 2G, 2H), suggesting that the succinylation of Fkbp3 at other sites also contributes to AFB1 biosynthesis in A. flavus. The above results indicated that the K19 site likely plays a more prominent role in Fkbp3 succinylation compared to other sites.

Added in the revised version: The Sch9 kinase is involved in aflatoxin biosynthesis and the HOG pathway

Added in the revised version: Sch9, a Ser/Thr kinase belonging to the AGC family, is directly phosphorylated by TORC1. The phosphorylation of Sch9 is inhibited by rapamycin, as well as by carbon or nitrogen starvation [22]. The putative sch9 gene (AFLA_127440) in A. flavus encodes a protein consisting of 706 amino acids, exhibiting 80.31% identity to the Sch9 protein in A. nidulans. To investigate the role of Sch9 in A. flavus, we attempted to create a sch9 deletion mutant using a homologous recombination strategy. The Δsch9 strain exhibited typical vegetative growth and conidiation similar to the wild-type strain (Fig. 3B). TLC assay and quantitative analysis showed a significantly decreased aflatoxin production in the Δsch9 strain compared to the wild-type strain (Fig. 3C, 3G). These results indicated that Sch9 regulates aflatoxin biosynthesis in A. flavus.

Removed in the revised version: The Tor kinase plays a critical role in A. flavus

Removed in the revised version: In S. cerevisiae, the Tor kinase can interact with a diverse range of proteins, resulting in the formation of complex structures known as TORC1 and TORC2 [44]. TORC1 can detect various signals, including nutrients, growth factors, and environmental stress. It plays a crucial role in regulating cellular processes such as gene transcription, protein translation, ribosome synthesis, and autophagy [45]. To identify the ortholog of the Tor protein in A. flavus, the S. cerevisiae Tor protein was utilized as search queries in the A. flavus genome database using the Basic Local Alignment Search Tool (https://blast.ncbi.nlm.nih.gov/Blast.cgi). The genome of A. flavus contains a solitary ortholog of the tor gene (AFLA_044350), which encodes a protein exhibiting 48.91% similarity to S. cerevisiae Tor2. We used the absolute quantitative PCR technique to identify the copy number of the tor gene at the genome level in A. flavus (Table S1). To elucidate the function of the Tor kinase in A. flavus, we employed gene deletion techniques to target the tor gene. While our attempt to obtain a null mutant was unsuccessful, we constructed a single copy deletion strain, tor+/−, through homologous recombination, which contains one copy of the tor gene (Table S1). Compared to the wild-type strain, the tor+/− strain exhibited a significantly slower growth rate on PDA (Fig. 3C). Additionally, the tor+/− strain produced fewer conidia and sclerotia compared to the wild-type strain (Fig. 3B, S4). The results from thin-layer chromatography (TLC) revealed a noteworthy decrease in AFB1 synthesis in the tor+/− strain compared to the wild-type strain (Fig. 3D, 3F). Interestingly, the tor+/− mutant exhibited a notable increase in sensitivity to rapamycin and decreased sensitivity to osmotic stress (Fig. 3C, 3E, S5). Furthermore, the qRT-PCR results demonstrated a significant downregulation in the expression levels of the genes associated with sexual sporulation and aflatoxin synthesis in the tor+/− mutant compared to the wild-type strain (Fig. S6). The transcriptional level of the tor gene in the tor+/− mutant was downregulated compared to the level in the wild-type strain (Fig. S6). These findings suggested that the Tor kinase plays a crucial role in various biological processes in A. flavus, including vegetative growth, asexual development, sclerotia formation, response to environmental stress, and biosynthesis of aflatoxins.

Removed in the revised version: The Tor kinase plays a vital role in A. flavus

Removed in the revised version: (A) The structure diagram of the Tor kinase.

Removed in the revised version: (B) Morphology of conidiophores in the WT, tor+/-, and torΔFRBstrains observed by microscope.

Removed in the revised version: (C) Phenotype of the WT, tor+/-, and torΔFRBstrains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Removed in the revised version: (D) TLC assay of AFB1 production by the WT, tor+/−, and torΔFRBstrains cultured in YES liquid medium at 29°C for 6 days.

Removed in the revised version: (E) The growth inhibition rate of the WT, tor+/-, and torΔFRBstrains under rapamycin stress.

Removed in the revised version: (F) AFB1 quantitative analysis of the WT, tor+/-, and torΔFRBstrains described in (D).

Removed in the revised version: The Tor protein is characterized by the presence of various conserved domains, including the DUF3385, FAT, FKBP12-rapamycin binding (FRB), kinase catalytic, and FATC domains, which are located at the N-terminus (Fig. 3A). To investigate the functional significance of the FRB domain within the Tor kinase, the FRB domain was eliminated using homologous recombination. Compared to the tor+/− strain, the torΔFRB mutant displayed comparable deficiencies in vegetative growth, conidiation, sclerotia formation, osmotic stress, and rapamycin sensitivity (Fig. 3B, 3C, S4, S5). Based on these observations, we speculated that the FRB domain plays a pivotal role in the modulation of the Tor kinase activity, thereby influencing fungal development and rapamycin resistance in A. flavus.

The Sch9 kinase participates in aflatoxin biosynthesis and the HOG pathway.

(A) The structure diagram of the Sch9 kinase.

(B) Phenotype of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc , sch9ΔS_TK_X , and sch9K340A strains grown on PDA medium amended with 100 ng/mL rapamycin at 37°C for 5 days.

(C) TLC assay of AFB1 production from the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc , sch9ΔS_TK_X , and sch9K340A strains cultured in YES liquid medium at 29°C for 6 days.

(D) The phenotypes of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc , sch9ΔS_TK_X , and sch9K340Astrains on YGT media amended with 1 M NaCl and 1 M KCl for 3 days.

(E) The phosphorylation levels of Hog1 in the WT and Δsch9 strains were determined with or without osmotic stress.

(F) The growth inhibition rate of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc , sch9ΔS_TK_X , and sch9K340A strains under rapamycin stress.

(G) Quantitative analysis of AFB1 as shown in (C).

(H) The growth inhibition rate of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc , sch9ΔS_TK_X , and sch9K340A strains under osmotic stress.

Removed in the revised version: TapA and TipA regulate sclerotia formation in A. flavus

Removed in the revised version: (A) Phenotype of the WT, tapA +/−, and ΔtipA strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Removed in the revised version: (B) Phenotypic characterization of the WT, tapA +/−, and ΔtipA strains grown on CM medium at 37°C for 7 days.

Removed in the revised version: (C) Morphology of conidiophores in the WT, tapA +/−, and ΔtipA strains observed by microscope.

Removed in the revised version: (D) TLC assay of AFB1 production by the WT, tapA +/−, and ΔtipA strains cultured in YES liquid medium at 29°C for 6 days.

Removed in the revised version: (E) The growth inhibition rate of the WT, tapA +/−, and ΔtipA strains under rapamycin stress.

Removed in the revised version: (F) Amount of sclerotia produced by the WT, tapA +/−, and ΔtipA strains.

Removed in the revised version: (G) AFB1 quantitative analysis of the WT, tapA +/−, and ΔtipA strains described in (D).

Removed in the revised version: The Sch9 kinase is involved in aflatoxin biosynthesis and the HOG pathway

Removed in the revised version: Sch9, a Ser/Thr kinase belonging to the AGC family, is directly phosphorylated by TORC1. The phosphorylation of Sch9 is inhibited by rapamycin, as well as by carbon or nitrogen starvation [22]. The putative sch9 gene (AFLA_127440) in A. flavus encodes a protein consisting of 706 amino acids, exhibiting 80.31% identity to the Sch9 protein in A. nidulans. To investigate the role of Sch9 in A. flavus, we attempted to create a sch9 deletion mutant using a homologous recombination strategy. The Δsch9 strain exhibited typical vegetative growth and conidiation similar to the wild-type strain (Fig. 4B). TLC assay and quantitative analysis showed a significantly decreased aflatoxin production in the Δsch9 strain compared to the wild-type strain (Fig. 4C, 4G). These results indicated that Sch9 regulates aflatoxin biosynthesis in A. flavus.

Earlier version: In yeast, Sch9 plays an important role in the transcriptional activation of osmostress inducible genes [46]. To characterize the roles of Sch9 in stress response, we examined the sensitivity of the sch9 deletion strain to various stresses including osmotic stress and calcium stress. Compared with the wild-type strain, the Δsch9 mutant significantly increased sensitivity to calcium stress (Fig. S7), and decreased sensitivity to osmotic stress induced by NaCl and KCl (Fig.4D, 4H). In A. flavus, the Sch9 kinase consists of several domains, including the protein kinase C conserved region 2 (C2), protein kinase (S_TKc), AGC-kinase C-terminal (S_TK_X) domain, and an ATP binding site (Fig. 4A). To explore the role of these domains and the ATP binding site in Sch9 kinase, we constructed deletion strains for both domains and a sch9K340A mutant. The results revealed that the sch9ΔS_TKc strain, sch9ΔS_TK_X strain, and sch9K340A mutant exhibited similar phenotypes in osmotic stress and AFB1 production as the Δsch9 mutant (Fig.4B, 4C, 4D). Therefore, we speculate that the protein kinase domain, kinase C-terminal domain, and the ATP binding site at K340 play a crucial role in modulating the impact of Sch9 on aflatoxin biosynthesis and stress response in A. flavus.

Revised version: In yeast, Sch9 plays an important role in the transcriptional activation of osmostress inducible genes [44]. To characterize the roles of Sch9 in stress response, we examined the sensitivity of the sch9 deletion strain to various stresses including osmotic stress and calcium stress. Compared with the wild-type strain, the Δsch9 mutant significantly increased sensitivity to calcium stress (Fig. S4A, S4B), and decreased sensitivity to osmotic stress induced by NaCl and KCl (Fig. 3D, 3H). In A. flavus, the Sch9 kinase consists of several domains, including the protein kinase C conserved region 2 (C2), protein kinase (S_TKc), AGC-kinase C-terminal (S_TK_X) domain, and an ATP binding site (Fig. 3A). To explore the role of these domains and the ATP binding site in Sch9 kinase, we constructed deletion strains for both domains and a sch9K340A mutant. The results revealed that the sch9ΔS_TKc strain, sch9ΔS_TK_X strain, and sch9K340A mutant exhibited similar phenotypes in osmotic stress and AFB1 production as the Δsch9 mutant (Fig. 3B, 3C, 3D). Therefore, we speculate that the protein kinase domain, kinase C-terminal domain, and the ATP binding site at K340 play a crucial role in modulating the impact of Sch9 on aflatoxin biosynthesis and stress response in A. flavus.

Earlier version: To investigate the potential function of Sch9 in the HOG and TOR pathways in A. flavus, we conducted a sensitivity analysis of the Δsch9 strain to rapamycin. Compared to the wild-type strain, the sch9 deletion strain exhibited a slightly higher sensitivity to rapamycin (Fig. 4B, 4F). The MAKP Hog1 is an element of the high osmolarity glycerol (HOG) pathway. To verify the connection between Sch9 and the HOG-MAPK pathway, the phosphorylation levels of Hog1 kinase were measured under osmotic stress (1M NaCl). Western blot analysis showed that the phosphorylation level of Hog1 increased significantly in the Δsch9 strain (Fig. 4E). This finding suggested that Sch9 plays a pivotal role in modulating the response of Hog1 to osmotic stress. Taken together, these results reflected that Sch9 regulates osmotic stress response via the HOG pathway in A. flavus.

Revised version: To investigate the potential function of Sch9 in the HOG and TOR pathways in A. flavus, we conducted a sensitivity analysis of the Δsch9 strain to rapamycin. Compared to the wild-type strain, the sch9 deletion strain exhibited a slightly higher sensitivity to rapamycin (Fig. 3B, 3F). The MAKP Hog1 is an element of the high osmolarity glycerol (HOG) pathway. To verify the connection between Sch9 and the HOG-MAPK pathway, the phosphorylation levels of Hog1 kinase were measured under osmotic stress (1M NaCl). Western blot analysis showed that the phosphorylation level of Hog1 increased significantly in the Δsch9 strain (Fig. 3E). This finding suggested that Sch9 plays a pivotal role in modulating the response of Hog1 to osmotic stress. Taken together, these results reflected that Sch9 regulates osmotic stress response via the HOG pathway in A. flavus.

Earlier version: TapA and TipA participate in regulation of sclerotia development and cell wall stress response

Revised version: TipA participate in regulation of sclerotia development and cell wall stress response

Removed in the revised version: Tap42, a protein that associates with phosphatase 2A, serves as the direct target of the Tor kinase in yeast. Tap42 plays a crucial role in various Tor functions, particularly in the regulation of transcriptional processes [47]. The genome of A. flavus contains a homolog of tap42 named tapA (AFLA_092770), which is predicted to encode a 355 amino-acid protein (sharing 31.9% identity with S. cerevisiae Tap42). To further determine the function of the TapA regulator in the TOR signaling pathway of A. flavus, we conducted targeted gene deletion experiments on tapA. We failed to obtain a null mutant, therefore we generated the tapA+/− strain, which contains one copy of the tapA gene (Table S2). Compared to the wild-type strain, the tapA+/− strain exhibited diminished hyphal growth and conidiation (Fig. 5A, 5C), indicating that TapA is important for vegetative growth and conidiation. The results of the rapamycin sensitivity test indicated that the tapA+/− strain displayed reduced sensitivity to rapamycin (Fig.5A, 5E). In addition, the TapA+/− strain was unable to generate any sclerotia (Fig. 5B, 5F), and almost lost its capability in AFB1 production compared to the wild-type strain (Fig. 5D, 5G), these findings indicated that TapA is essential for sclerotia formation and biosynthesis of AFB1. The tapA+/− mutant was more insensitive to cell wall stress, including CFW and CR (Fig. S8). Furthermore, the levels of nsdC, nsdD, and sclR expression were significantly decreased in tapA+/− strain compared to those in the wild-type strain (Fig. S9). The transcriptional level of the tapA gene in the tapA+/− strain was drastically downregulated compared to the wild-type strain (Fig. S9). These findings have demonstrated the significantly positive impact of TapA on the growth, sporulation, sclerotia formation, cell wall stress, and AF biosynthesis in A. flavus.

Earlier version: In S. cerevisiae, Tip41 was identified as a Tap42-interacting protein and has been measured to function as a negative regulator of Tap42, which downregulates TORC1 signaling via activation of PP2A phosphatase [48]. The putative orthologue of tip41 in A. flavus named tipA (AFLA_047310). The amino acid sequence of TipA exhibits a similarity of 40.49% to Tap41 in S. cerevisiae. To elucidate the function of the TOR signaling pathway protein TipA in A. flavus, we generated the tipA deletion mutants. The ΔtipA strain exhibited a slightly reduced number of conidia produced on PDA compared to the wild-type strain (Fig. 5A). Meanwhile, the ΔtipA strain exhibited decreased sensitivity to rapamycin (Fig. 5A). The ΔtipA strain was unable to produce any sclerotia (Fig. 5B). In addition, the ΔtipA strain was more sensitive to cell wall and cell membrane stress, including CR and SDS (Fig. S8). Additionally, the expression levels of nsdC and nsdD exhibited a decrease in ΔtipA strain compared to the wild-type strain (Fig. S9). These results indicated that TipA is crucial for the sclerotial formation and cell wall stress in A. flavus.

Revised version: In S. cerevisiae, Tip41 was identified as a Tap42-interacting protein and has been measured to function as a negative regulator of Tap42, which downregulates TORC1 signaling via activation of PP2A phosphatase [45]. The putative orthologue of tip41 in A. flavus is named tipA (AFLA_047310). The amino acid sequence of TipA exhibits a similarity of 40.49% to Tip41 in S. cerevisiae. To elucidate the function of the TOR signaling pathway protein TipA in A. flavus, we generated the tipA deletion mutants. The ΔtipA strain exhibited a slightly reduced number of conidia produced on PDA compared to the wild-type strain (Fig. 4A, 4C). Meanwhile, the ΔtipA strain exhibited decreased sensitivity to rapamycin (Fig. 4A, 4E). The ΔtipA strain was unable to produce any sclerotia (Fig. 4B, 4F). Whereas ΔtipA strain do not affect aflatoxin production(Fig. 4D, 4G). In addition, the ΔtipA strain was more sensitive to the cell wall and cell membrane stress, including CR and SDS (Fig. 4H, 4I). Additionally, the expression levels of nsdC and nsdD exhibited a decrease in ΔtipA strain compared to the wild-type strain (Fig. 4J). These results indicated that TipA is crucial for the sclerotial formation and cell wall stress in A. flavus.

Added in the revised version: TipA regulates clerotia development and cell wall stress in A. flavus.

Added in the revised version: (A) Phenotype of the WT and ΔtipA strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Added in the revised version: (B) Phenotypic characterization of the WT and ΔtipA strains grown on CM medium at 37°C for 7 days.

Added in the revised version: (C) Morphology of conidiophores in the WT and ΔtipA strains observed by microscope.

Added in the revised version: (D) TLC assay of AFB1 production by the WT and ΔtipA strains cultured in YES liquid medium at 29°C for 6 days.

Added in the revised version: (E) The growth inhibition rate of the WT and ΔtipA strains under rapamycin stress.

Added in the revised version: (F) Amount of sclerotia produced by the WT and ΔtipA strains.

Added in the revised version: (G) AFB1 quantitative analysis of the WT and ΔtipA strains described in (D).

Added in the revised version: (H) Colony morphology of the WT and ΔtipA strains grown on PDA media supplemented with 200 µg/mL CR, 200µg/mL CFW, and 100 µg/mLSDS at 37°C for 5 days.

Added in the revised version: (I) The growth inhibition rate of the WT and ΔtipA strains under cell wall and cell membrane stress.

Added in the revised version: (J) Relative expression levels of sclerotia formation genes in the WT and ΔtipA strains.

Phosphatase SitA and Ppg1 are involved in growth, conidiation, and sclerotial formation

Earlier version: In S. cerevisiae, Tap42 has been identified to interact with the catalytic subunit of type 2A protein phosphatases, including Pph3, Pph21, and Pph22, as well as the type 2A-like phosphatases Sit4 and Ppg1[19]. Phosphatase sitA (AFLA_108450) and ppg1 (AFLA_132610) were identified in A. flavus, encoding putative orthologs of sit4 and ppg1 in S. cerevisiae. To determine the biological functions of phosphatase SitA and Ppg1 in A. flavus, we disrupted the sitA and ppg1 by homologous recombination strategy. To further investigate the relationship between phosphatase SitA and Ppg1, we additionally generated complemented and double deletion mutants. The vegetative development phenotype analysis showed that the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains grew significantly more slowly than the wild-type strain on PDA (Fig. 6A, 6D). Additionally, the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains were more sensitive to rapamycin (Fig. 6A, 6E). In addition, the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains formed lower number of conidiophores (Fig. 6A). To gain further insight into the role of SitA and Ppg1 in the process of conidiation, qRT-PCR was employed to determine the transcript levels of conidial synthesis-related genes (abaA and brlA). The expression levels of the abaA and brlA genes were significantly downregulated in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants when compared to the wild-type and complemented strains (Fig. 6G). The analysis of sclerotial formation revealed that the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants exhibited a complete inability to produce any sclerotia (Fig. 6B). To further prove these findings, we conducted qRT-PCR to determine the expression levels of three regulators (nsdC, nsdD, and nclR) of sclerotia development. Our findings indicated that the expression levels of nsdC, nsdD, and sclR were all significantly reduced in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants (Fig. 6H), which is consistent with the observed decrease in sclerotia production in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. These results showed that SitA and Ppg1 are crucial for hyphal development, conidiation, and sclerotia formation in A. flavus.

Revised version: In S. cerevisiae, Tap42 has been identified to interact with the catalytic subunit of type 2A protein phosphatases, including Pph3, Pph21, and Pph22, as well as the type 2A-like phosphatases Sit4 and Ppg1[19]. Phosphatase sitA (AFLA_108450) and ppg1 (AFLA_132610) were identified in A. flavus, encoding putative orthologs of sit4 and ppg1 in S. cerevisiae. To determine the biological functions of phosphatase SitA and Ppg1 in A. flavus, we disrupted the sitA and ppg1 by homologous recombination strategy. To further investigate the relationship between phosphatase SitA and Ppg1, we additionally generated complemented and double deletion mutants. The vegetative development phenotype analysis showed that the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains grew significantly more slowly than the wild-type strain on PDA (Fig. 5A, 5D). Additionally, the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains were more sensitive to rapamycin (Fig. 5A, 5E). In addition, the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains formed lower number of conidiophores (Fig. 5A ). To gain further insight into the role of SitA and Ppg1 in the process of conidiation, qRT-PCR was employed to determine the transcript levels of conidiation-related genes (abaA and brlA). The expression levels of the abaA and brlA genes were significantly downregulated in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants when compared to the wild-type and complemented strains (Fig. 5G). The analysis of sclerotial formation revealed that the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants exhibited a complete inability to produce any sclerotia (Fig. 5B, 5E). To further prove these findings, we conducted qRT-PCR to determine the expression levels of three regulators (nsdC, nsdD, and sclR) of sclerotia development. Our findings indicated that the expression levels of nsdC, nsdD, and sclR were all significantly reduced in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants (Fig. 5H), which is consistent with the observed decrease in sclerotia production in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. These results showed that SitA and Ppg1 are crucial for hyphal development, conidiation, and sclerotia formation in A. flavus.

The impact of the phosphatases SitA and Ppg1 on the growth, conidiation, and aflatoxin biosynthesis in A. flavus.

(A) Colony morphology of the WT, single knockout, and double knockout strains cultured on PDA medium amended with 100 ng/mL rapamycin at 37℃ for 5 days.

Earlier version: (B) Colony morphology of the WT, single knockout, and double knockout strains cultured on WKM medium at 37℃ for 7 days.

Revised version: (B) Colony morphology of the WT, single knockout, and double knockout strains cultured on CM medium at 37℃ for 7 days.

(C) TLC analysis of AFB1 production from the WT, single knockout, and double knockout strains cultured in YES liquid medium at 29℃ for 6 days.

(D) Growth diameter of the WT, single knockout, and double knockout strains on PDA media.

(E) The growth inhibition rate of the WT, single knockout, and double knockout strains under rapamycin stress.

(F) AFB1 quantitative analysis of the WT, single knockout, and double knockout strains.

(G) Relative expression levels of conidia synthesis genes in the WT, single knockout, and double knockout strains.

(H) Relative expression levels of sclerotia synthesis genes in the WT, single knockout, and double knockout strains.

(I) Relative expression levels of aflatoxin biosynthesis genes in the WT, single knockout, and double knockout strains.

SitA and Ppg1 play crucial roles in aflatoxin biosynthesis and pathogenicity

Earlier version: We assayed AFB1 biosynthesis in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. TLC analyses revealed that the ΔsitA and ΔsitA/Δppg1 strains exhibited a decrease in AFB1 production, whereas the Δppg1 strain showed a significant increase in AFB1 production compared to the wild-type and complemented strains (Fig. 6C, 6F). The expression levels of genes related to aflatoxin synthesis, including structural genes (aflO and aflP) and regulatory genes (aflR and aflS), were determined by qRT-PCR in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. The results revealed a significant downregulation of these genes in the ΔsitA and ΔsitA/Δppg1 mutants, while upregulated in the Δppg1 mutants, compared to the wild-type and complemented strains (Fig. 6I). These results indicated that SitA and Ppg1 play different roles in regulating aflatoxin biosynthesis in A. flavus.

Revised version: We assayed AFB1 biosynthesis in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. TLC analyses revealed that the ΔsitA and ΔsitA/Δppg1 strains exhibited a decrease in AFB1 production, whereas the Δppg1 strain showed a significant increase in AFB1 production compared to the wild-type and complemented strains (Fig. 5C, 5F). The expression levels of genes related to aflatoxin synthesis, including structural genes (aflO and aflP) and regulatory genes (aflR and aflS), were determined by qRT-PCR in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. The results revealed a significant downregulation of these genes in the ΔsitA and ΔsitA/Δppg1 mutants, while upregulated in the Δppg1 mutants, compared to the wild-type and complemented strains (Fig. 5I). These results indicated that SitA and Ppg1 play different roles in regulating aflatoxin biosynthesis in A. flavus.

Earlier version: We conducted pathogenicity tests on peanuts and maize to examine the pathogenicity of the phosphatases SitA and Ppg1 on crops. The results showed that the wild-type and complemented strains demonstrated complete virulence on all peanut and maize seeds. In contrast, the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants displayed compromised colonization on peanut and maize seeds (Fig. 7A). In addition, the number of conidial productions in these mutants on the infected peanut seeds also dramatically decreased compared to those of the wild-type and complemented strains (Fig. 7C). The thin-layer chromatography analyses revealed that the Δppg1 mutant exhibited increased production of AFB1 in peanuts and maize. Conversely, the levels of aflatoxin produced by the ΔsitA and ΔsitA/Δppg1 mutants on peanut and maize were significantly reduced compared to the the wild-type and complemented strains (Fig. 7B, 7D). This observation aligns with the above results of aflatoxin biosynthesis from all mutant strains in the YES medium. All these data showed that SitA and Ppg1 are crucial for crop seeds’ pathogenicity.

Revised version: We conducted pathogenicity tests on peanuts and maize to examine the pathogenicity of the phosphatases SitA and Ppg1 on crops. The results showed that the wild-type and complemented strains demonstrated complete virulence on all peanut and maize seeds. In contrast, the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants displayed compromised colonization on peanut and maize seeds (Fig. 6A). In addition, the number of conidial productions in these mutants on the infected peanut seeds also dramatically decreased compared to those of the wild-type and complemented strains (Fig. 6C). The thin-layer chromatography analyses revealed that the Δppg1 mutant exhibited increased production of AFB1 in peanuts and maize. Conversely, the levels of aflatoxin produced by the ΔsitA and ΔsitA/Δppg1 mutants on peanut and maize were significantly reduced compared to the wild-type and complemented strains (Fig. 6B, 6D). This observation aligns with the above results of aflatoxin biosynthesis from all mutant strains in the YES medium. All these data showed that SitA and Ppg1 are crucial for crop seeds’ pathogenicity.

Pathogenicity analysis of the phosphatase SitA and Ppg1 in A. flavus.

(A) Phenotypes of peanut and maize infected by the WT, single knockout, and double knockout strains.

(B) TLC was used to detect the AFB1 from peanut and maize infected by the WT, single knockout, and double knockout strains.

(C) Quantitative analysis of conidia production from peanut and maize infected by the WT, single knockout, and double knockout strains.

(D) Quantitative analysis of AFB1 from peanut and maize infected by the WT, single knockout, and double knockout strains.

SitA and Ppg1 are involved in regulating cell wall integrity

Earlier version: To explore the role of phosphatases SitA and Ppg1 in response to cell wall stress, all strains were cultivated on PDA amended with CFW, CR, and SDS. We found that the relative growth inhibition of ΔsitA, Δppg1 and ΔsitA/Δppg1 mutants induced by cell wall stress were significantly increased compared to the wild-type and complemented strains (Fig. 8A, 8B). Further investigations using qRT-PCR revealed a notable reduction in the expression levels of genes related to the synthesis of cell walls, specifically chitin synthase genes chsA, chsD, chsE, and β-1,6, in the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains (Fig. 8C). Western blot analysis also revealed that the phosphorylation level of Slt2 was significantly elevated in the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains under cell wall stress (200 μg/ml CFW) (Fig. 8D). In conclusion, the phosphatases SitA and Ppg1 play a crucial role in regulating cell wall integrity.

Revised version: To explore the role of phosphatases SitA and Ppg1 in response to cell wall stress, all strains were cultivated on PDA amended with CFW, CR, and SDS. We found that the relative growth inhibition of ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants induced by cell wall stress was significantly increased compared to the wild-type and complemented strains (Fig. 7A, 7B). Further investigations using qRT-PCR revealed a notable reduction in the expression levels of genes related to the synthesis of cell walls, specifically chitin synthase genes chsA, chsD, and chsE in the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains (Fig. 7C). Western blot analysis also revealed that the phosphorylation level of Slt2 was significantly elevated in the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains under cell wall stress (200 μg/ml CFW) (Fig. 7D). In conclusion, the phosphatases SitA and Ppg1 play a crucial role in regulating cell wall integrity.

Sensitivity of phosphatase SitA and Ppg1 to cell wall damaging agents in A. flavus.

(A) Morphology of the WT, single knockout, and double knockout strains grown on YES media supplemented with 200 µg/mL CR, 200 µg/mL CFW, or 100 µg/mL SDS at 37°C for 5 days.

(B) The growth inhibition rate of the WT, single knockout, and double knockout strains under cell wall and cell membrane stress.

Earlier version: (C) Relative expression levels of cell wall synthesis genes (chsA, chsD, chsE and β-1,6) in the WT, single knockout, and double knockout strains.

Revised version: (C) Relative expression levels of cell wall synthesis genes (chsA, chsD and chsE) in the WT, single knockout, and double knockout strains.

(D) The phosphorylation level of Slt2 in the WT, single knockout, and double knockout strains was detected with or without cell wall stress.

Earlier version: The phosphatase complex Nem1/Spo7 are important for growth, aflatoxin biosynthesis, and LD biogenesis

Revised version: The phosphatase complex Nem1/Spo7 is important for growth, aflatoxin biosynthesis, and LD biogenesis

Earlier version: Lipid droplets (LDs) are highly dynamic spherical organelles, which appear in the cytosol of most eukaryotic cell types[49]. Rapamycin treatment led to the accumulation of lipid droplets in various filamentous fungi [24]. The FgTOR pathway regulates the phosphorylation status of FgNem1 mainly via the FgPpg1/Sit4-FgCak1/other kinase cascade in F. graminearum [24]. In this study, we observed that the treatment of hyphae with rapamycin resulted in the accumulation of LD biogenesis (Fig. 9A). This finding suggested that the TOR pathway may regulate the LD biogenesis. However, the underlying regulatory mechanism of LD biosynthesis remain largely unknown in A. flavus. To investigate the regulatory mechanism of the TOR pathway in A. flavus that controls LD biogenesis, we examined LD accumulation in Δsch9, ΔsitA, and Δppg1 mutants. BODIPY staining assays revealed that LD biogenesis was not observed in the ΔsitA and Δppg1 strains upon treatment with rapamycin, in contrast to the wild-type and Δsch9 strains (Fig. 9A, 9B, 9C, 9D). These results indicated that LD biogenesis induced by rapamycin is predominantly reliant on the phosphatases SitA and Ppg1 in A. flavus. In F. graminearum, the phosphatase FgNem1/FgSpo7 complex was involved in the induction of LD biogenesis by rapamycin. Phosphatase nem1 (AFLA_002649) and spo7 (AFLA_028590) were identified in A. flavus, which encode putative orthologs of nem1 and spo7 in F. graminearum. To examine the role of the nem1 and spo7 genes in A. flavus, we generated Δnem1 and Δspo7 mutants using homologous recombination. The Δnem1 and Δspo7 mutants displayed a notable reduction in both vegetative growth and conidiation compared to the wild-type strain (Fig. 10A, 10C, 10D). TLC assay revealed a significantly decreased aflatoxin production in the Δnem1 and Δspo7 compared to the wild-type strain (Fig. 10B). In addition, BODIPY staining showed no visible lipid droplets in the hyphae of Δnem1 and Δspo7 after rapamycin treatment (Fig. 10E). Taken together, these results indicated that Nem1/Spo7 are involved in vegetative growth, conidiation, aflatoxin production, and LD biogenesis.

Revised version: Lipid droplets (LDs) are highly dynamic spherical organelles, which appear in the cytosol of most eukaryotic cell types[46]. Rapamycin treatment led to the accumulation of lipid droplets in various filamentous fungi [24]. The FgTOR pathway regulates the phosphorylation status of FgNem1 mainly via the FgPpg1/Sit4-FgCak1/other kinase cascade in F. graminearum [24]. In this study, we observed that the treatment of hyphae with rapamycin resulted in the accumulation of LD biogenesis (Fig. 8A). This finding suggested that the TOR pathway may regulate the LD biogenesis. However, the underlying regulatory mechanism of LD biosynthesis remain largely unknown in A. flavus. To investigate the regulatory mechanism of the TOR pathway in A. flavus that controls LD biogenesis, we examined LD accumulation in Δsch9, ΔsitA, and Δppg1 mutants. BODIPY staining assays revealed that LD biogenesis was not observed in the ΔsitA and Δppg1 strains upon treatment with rapamycin, in contrast to the wild-type and Δsch9 strains (Fig. 8A, 8B, 8C, 8D). These results indicated that LD biogenesis induced by rapamycin is predominantly reliant on the phosphatases SitA and Ppg1 in A. flavus. In F. graminearum, the phosphatase FgNem1/FgSpo7 complex was involved in the induction of LD biogenesis by rapamycin[24]. Phosphatase nem1 (AFLA_002649) and spo7 (AFLA_028590) were identified in A. flavus, which encode putative orthologs of nem1 and spo7 in F. graminearum. To examine the role of the nem1 and spo7 genes in A. flavus, we generated Δnem1 and Δspo7 mutants using homologous recombination.

Phosphatase SitA and Ppg1 are involved in regulating lipid droplet biogenesis.

(A) The intracellular lipid droplets were stained by boron dipyrromethene difluoride (BODIPY) in hyphae of the WT strain and observed under fluorescence microscopy. The phenotype of lipid droplets accumulation in the mycelia of the WT strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

(B) Phenotype of lipid droplets accumulation in the mycelia of the Δsch9 strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

(C) Phenotype of lipid droplets accumulation in the mycelia of the ΔsitA strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

(D) Phenotype of lipid droplets accumulation in the mycelia of the Δppg1 strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

Added in the revised version: The Δnem1 and Δspo7 mutants displayed a notable reduction in both vegetative growth and conidiation compared to the wild-type strain (Fig. 9A, 9C, 9D). TLC assay revealed a significantly decreased aflatoxin production in the Δnem1 and Δspo7 compared to the wild-type strain (Fig. 9B). In addition, BODIPY staining showed no visible lipid droplets in the hyphae of Δnem1 and Δspo7 after rapamycin treatment (Fig. 9E). Taken together, these results indicated that Nem1/Spo7 are involved in vegetative growth, conidiation, aflatoxin production, and LD biogenesis.

Phosphatase complex Nem1/Spo7 play significant role in growth, conidiation, aflatoxin, and lipid droplet biogenesis.

(A) Phenotype of the WT, Δnem1, and Δspo7 strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

(B) TLC assay of AFB1 production from the WT, Δnem1, and Δspo7 strains cultured in YES liquid medium at 29°C for 6 days.

(C) Growth diameter of the WT, Δnem1, and Δspo7 strains on PDA media.

(D) Statistical analysis of the sporulation by the WT, Δnem1, and Δspo7 strains.

(E) Phenotype of lipid droplets accumulation in the mycelia of the WT, Δnem1, and Δspo7 strains treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

Discussion

Earlier version: Aflatoxins (AFs) are a class of toxic secondary metabolites synthesized by A. flavus, including aflatoxin B1, B2, G1, and G2, among which AFB1 has been regarded as the most pathogenic and toxic aflatoxin[50]. Therefore, understanding the regulatory mechanisms involved in fungal secondary metabolites biosynthesis would be instrumental for effectively addressing aflatoxin control in A. flavus. In this study, we demonstrated that the target of rapamycin (TOR) signaling pathway plays a pivotal role in regulating growth, sporulation, sclerotia formation, aflatoxin production, and various stress response in A. flavus. Through conducting functional studies on genes within the TOR signaling pathway, we explored the conservation and complexity exhibited by the TOR signaling pathway in A. flavus. In addition, we constructed a crosstalk network between the TOR and other signaling pathways (HOG, CWI) and analyzed the specific regulatory process. We proposed that the TOR signaling pathway interacts with multiple signaling pathways, including HOG and CWI pathway, to modulate the cellular responses to diverse environmental stresses (Fig. 11).

Revised version: Aflatoxins (AFs) are a class of toxic secondary metabolites synthesized by A. flavus, including aflatoxin B1, B2, G1, and G2, among which AFB1 has been regarded as the most pathogenic and toxic aflatoxin[47]. Therefore, understanding the regulatory mechanisms involved in fungal secondary metabolites biosynthesis would be instrumental for effectively addressing aflatoxin control in A. flavus. In this study, we demonstrated that the target of rapamycin (TOR) signaling pathway plays a pivotal role in regulating growth, sporulation, sclerotia formation, aflatoxin production, and various stress responses in A. flavus. Through conducting functional studies on genes within the TOR signaling pathway, we explored the conservation and complexity exhibited by the TOR signaling pathway in A. flavus. In addition, we constructed a crosstalk network between the TOR and other signaling pathways (HOG, CWI) and analyzed the specific regulatory process. We proposed that the TOR signaling pathway interacts with multiple signaling pathways, including HOG and CWI pathways, to modulate the cellular responses to diverse environmental stresses (Fig. 10).

The proposed model of the target of rapamycin (TOR) pathway in A. flavus.

Earlier version: Based on the above results, we propose a hypothesis regarding the TOR signaling pathway model in A. flavus. Rapamycin forms a complex with Fkbp3, and this particular complex can bind to the Tor kinase, thereby impeding its regular functionality. Sch9, functioning as a downstream element of the Tor kinase, regulates aflatoxin biosynthesis and HOG signaling pathway. As another important target of the Tor kinase, TapA-phosphatase complex is involved in the regulation of growth, sporulation, sclerotia, aflatoxin production, CWI signaling pathway, lipid droplet synthesis, and other processes. In conclusion, the TOR signaling pathway plays a crucial role in various aspects of A. flavus, including vegetative development, stress response, aflatoxin biosynthesis, and pathogenicity.

Revised version: Based on the above results, we propose a hypothesis regarding the TOR signaling pathway model in A. flavus. Rapamycin forms a complex with Fkbp3, and this particular complex can bind to the TorA kinase, thereby impeding its regular functionality. Sch9, functioning as a downstream element of the TorA kinase, regulates aflatoxin biosynthesis and HOG signaling pathway. As another important target of the TorA kinase, TapA-phosphatase complex is involved in the regulation of growth, sporulation, sclerotia, aflatoxin production, CWI signaling pathway, lipid droplet synthesis, and other processes. In conclusion, the TOR signaling pathway plays a crucial role in various aspects of A. flavus, including vegetative development, stress response, aflatoxin biosynthesis, and pathogenicity.

Earlier version: We have observed that rapamycin has a significant impact on the growth and sporulation of A. flavus (Fig. 1A). Additionally, it exhibits a strong inhibitory effect on sclerotia production and aflatoxin biosynthesis (Fig. 1B, 1C). The inhibitory effect has been documented in various fungal species, including Candida albicans [51], Cryptococcus neoformans [51], Podospora anserine [27], A. nidulans [28], Magnaporthe oryzae [39], and F. graminearum [23]. This suggested that the TOR signaling pathway may have a conserved function in filamentous fungi. The number of FK506-binding proteins varies among different species: there are four Fkbps protein in S. cerevisiae and three Fkbps protein in S. pombe [52]. In S. cerevisiae, rapamycin can bind to Fkbp protein, resulting in the irreversible inhibition of the G1 phase of the cell cycle and the regulation of cell growth [53]. In A. nidulans, the fkbp12 gene homologue fprA displays considerable sequence similarity with homologues identified in the S. cerevisiae [28]. In this study, based on the homologous alignment of the protein sequence of FprA, it was observed that four fkbp genes present in A. flavus. Among these, deletion of fkbp3 exhibited enhanced tolerance to rapamycin (Fig. 2A). So we speculated that rapamycin forms a complex with Fkbp3 in A. flavus, subsequently targeting downstream target genes in order to exert its biological function. We found that Fkbp3 is involved in sclerotia and aflatoxin biosynthesis in A. flavus (Fig. 2A, S3), and subsequent analysis revealed that the underlying mechanism may be associated with succinylation modification. The succinylation sites of Fkbp3 were identified through the utilization of succinylation proteomics data. By constructing point mutant strains and conducting phenotype experiments, we found that the succinylation of Fkbp3 at residue K19 is involved in rapamycin resistance and aflatoxin biosynthesis (Fig. 2E, 2G). Previous studies have demonstrated that Fkbp12 has been identified as the receptor for FK506 and rapamycin, which inhibits calcineurin and target of rapamycin complex 1 (TORC1), respectively [53]. The deletion of Fkbp12 confers the resistance to rapamycin in C. albicans [51] and C. neoformans [51]. The entomopathogenic fungus Beauveria bassiana contains three putative fkbp genes, and disruption of Bbfkbp12 significantly increases the resistance to rapamycin and FK506 [54]. The plant pathogenic fungus F. graminearum harbors three Fkbps proteins, among which deletion of Fgfkbp12 leads to resistance to rapamycin [23]. Thus, we propose that these Fkbps may exhibit analogous functionalities in conferring resistance to rapamycin. In Botrytis cinerea, deletion of Bcfkbp12 cause a reduction of the virulence of strain T4, while knockout of the fkbp12 gene did not affect the pathogenic development of the strain B05.10 [32]. These findings demonstrated that the fkbp12 homologous genes exhibited conservation in terms of rapamycin resistance. However, there is significant variation in other biological functions among fungal species, which may be attributed to species or strain evolutionary divergence.

Revised version: We have observed that rapamycin has a significant impact on the growth and sporulation of A. flavus (Fig. 1A). Additionally, it exhibits a strong inhibitory effect on sclerotia production and aflatoxin biosynthesis (Fig. 1B, 1C). The inhibitory effect has been documented in various fungal species, including Candida albicans [48], Cryptococcus neoformans [48], Podospora anserine [27], A. nidulans [28], Magnaporthe oryzae [39], and F. graminearum [23]. This suggested that the TOR signaling pathway may have a conserved function in filamentous fungi. The number of FK506-binding proteins varies among different species: there are four Fkbps proteins in S. cerevisiae and three Fkbps proteins in S. pombe [49]. In S. cerevisiae, rapamycin can bind to Fkbp protein, resulting in the irreversible inhibition of the G1 phase of the cell cycle and the regulation of cell growth [50]. In A. nidulans, the fkbp12 gene homologue fprA displays considerable sequence similarity with homologues identified in the S. cerevisiae [28]. In this study, based on the homologous alignment of the protein sequence of FprA, it was observed that four fkbp genes are present in A. flavus. Among these, deletion of fkbp3 exhibited enhanced tolerance to rapamycin (Fig. 2A). So we speculated that rapamycin forms a complex with Fkbp3 in A. flavus, subsequently targeting downstream target genes to exert its biological function. We found that Fkbp3 is involved in sclerotia formation and aflatoxin biosynthesis in A. flavus (Fig. 2A, S3), and subsequent analysis revealed that the underlying mechanism may be associated with succinylation modification. The succinylation sites of Fkbp3 were identified through the utilization of succinylation proteomics data. By constructing point mutant strains and conducting phenotype experiments, we found that the succinylation of Fkbp3 at residue K19 is involved in rapamycin resistance and aflatoxin biosynthesis (Fig. 2E, 2G). Previous studies have demonstrated that Fkbp12 has been identified as the receptor for FK506 and rapamycin, which inhibits calcineurin and target rapamycin complex 1 (TORC1), respectively [50]. The deletion of Fkbp12 confers the resistance to rapamycin in C. albicans [48] and C. neoformans [48]. The entomopathogenic fungus Beauveria bassiana contains three putative fkbp genes, and disruption of Bbfkbp12 significantly increases the resistance to rapamycin and FK506 [51]. The plant pathogenic fungus F. graminearum harbors three Fkbps proteins, among which deletion of Fgfkbp12 leads to resistance to rapamycin [23]. Thus, we propose that these Fkbps may exhibit analogous functionalities in conferring resistance to rapamycin. In Botrytis cinerea, deletion of Bcfkbp12 caused a reduction of the virulence of strain T4, while knockout of the fkbp12 gene did not affect the pathogenic development of strain B05.10 [32]. These findings demonstrated that the fkbp12 homologous genes exhibited conservation in terms of rapamycin resistance. However, there is significant variation in other biological functions among fungal species, which may be attributed to species or strain evolutionary divergence.

Earlier version: The Tor kinase functions as a fundamental component of the TOR signaling pathway, which exhibits evolutionary conservation across fungi, plants, and mammals [55]. TOR assembles into two distinct multi-protein complexes known as TORC1 and TORC2. A characteristic that distinguishes TORC1 is its interaction with KOG1 (Raptor in mammals) and its sensitivity to the rapamycin-FKBP12 complex [56]. Almost all other eukaryotes, including plants, worms, flies, and mammals, have a single tor gene, whereas S. cerevisiae has two tor genes [45]. The Tor kinase is widely conserved both in terms of its structural characteristics and its role as the target of Fkbp-rapamycin. Typically, the Tor kinase exhibits comparable structural characteristics, including the HEAT-like, FAT, FRB, kinase, and FATC domains. The FRB domain is the binding region of the Fkbp-rapamycin complex, and rapamycin resistance is usually caused by the disruption of the FRB domain [45]. The comparison of homology and the analysis of phylogenetics revealed that there is only one predicted tor gene in A. flavus. To ascertain its functionality, we generated the tor single-copy deletion mutant by homologous recombination. We discovered that the Tor kinase is involved in the vegetative growth and asexual development of A. flavus (Fig. 3B, 3C). In addition, the tor+/− mutant shows a high sensitivity to rapamycin (Fig. 3C). Our findings have demonstrated the significant involvement of Tor kinase and its FRB domain in the processes of morphogenesis, aflatoxin biosynthesis, and multiple stress responses in A. flavus. We speculate the common expectation that the deletion of the tor and FRB domain should result in insensitivity and resistance to rapamycin, as it disrupts the binding site for Fkbp-rapamycin. However, we observed that the tor+/− and FRB domain-deleted mutant were more sensitive to rapamycin. This intriguing result suggests that there are additional factors or complexities involved in TOR signaling pathway regulation in A. flavus. We hypothesize that this result is related to the double copy of the tor gene. Previous findings demonstrated that the Tor kinase could function as a molecular switch connecting an iron cue to defend against pathogen infection in C. elegans [57]. In A. fumigatus, the Tor kinase represents a central regulatory node controlling genes and proteins involved in nutrient sensing, stress response, cell cycle progression, protein biosynthesis, and degradation, as well as adaptation to low iron conditions[58]. In summary, the Tor kinase acts as a central regulator in the TOR signaling pathway, playing a crucial role in the regulation of the cell cycle, protein synthesis, and cellular energy metabolism.

Revised version: The Tor kinase functions as a fundamental component of the TOR signaling pathway, which exhibits evolutionary conservation across fungi, plants, and mammals [52]. In S. cerevisiae, the Tor kinase can interact with a diverse range of proteins, resulting in the formation of complex structures known as TORC1 and TORC2 [53]. TORC1 can detect various signals, including nutrients, growth factors, and environmental stress. It plays a crucial role in regulating cellular processes such as gene transcription, protein translation, ribosome synthesis, and autophagy [54]. A characteristic that distinguishes TORC1 is its interaction with KOG1 (Raptor in mammals) and its sensitivity to the rapamycin-FKBP12 complex [55]. Almost all other eukaryotes, including plants, worms, flies, and mammals, have a single tor gene, whereas S. cerevisiae has two tor genes [54]. The Tor kinase is widely conserved both in terms of its structural characteristics and its role as the target of Fkbp-rapamycin. Typically, the Tor kinase exhibits similar structural characteristics, including the HEAT-like, FAT, FRB, kinase, and FATC domains. The FRB domain is the binding region of the Fkbp-rapamycin complex, and rapamycin resistance is usually caused by the disruption of the FRB domain [54]. The comparison of homology and the analysis of phylogenetics revealed that there is only one predicted torA gene in A. flavus. To ascertain its functionality, several ectopic transformants were obtained through multiple protoplast transformations, however, no successful knock-out transformants were obtained. We speculated that the deletion of the torA gene may be lethal in A. flavus. This observation aligns with the situation observed in F. graminearum [23]. Previous findings demonstrated that the Tor kinase could function as a molecular switch connecting an iron cue to defend against pathogen infection in C. elegans [56]. In A. fumigatus, the Tor kinase represents a central regulatory node controlling genes and proteins involved in nutrient sensing, stress response, cell cycle progression, protein biosynthesis, and degradation, as well as adaptation to low iron conditions[57]. In summary, the Tor kinase acts as a central regulator in the TOR signaling pathway, playing a crucial role in the regulation of the cell cycle, protein synthesis, and cellular energy metabolism.

Earlier version: In S. cerevisiae, the protein kinase Sch9 is a major downstream effector of TORC1, and this kinase is known to have crucial functions in stress resistance, longevity, and nutrient sensing [46]. In addition, the TORC1-Sch9 pathway acts as a crucial mediator of chronological lifespan [59]. In filamentous fungi, an intricate network of interactions exists between the TOR and HOG pathways. In F. graminearum [38], FgSch9 and FgHog1 mutants displayed heightened susceptibility to osmotic and oxidative stresses. The phosphorylation level of FgHog1 increased significantly in FgSch9 mutant. Additionally, the affinity capture-MS assay showed that the Tor kinase was among the proteins co-purifying with FgSch9. These results suggest that FgSch9 serves as an intermediary for the TOR and HOG signaling pathways [38]. Furthermore, FgSch9 plays a crucial role in the regulation of various biological processes such as hyphal differentiation, asexual development, virulence, and DON biosynthesis[38]. In A. fumigatus, the SchA mutant was sensitive to rapamycin, high concentrations of calcium and hyperosmotic stress, and SchA was involved in iron metabolism. The SchA mutant exhibited enhanced phosphorylation of SakA in response to osmotic stress [51]. In this study, we discovered that Sch9 is involved in regulating aflatoxin biosynthesis (Fig.4C), and responses to osmotic stress and rapamycin stress (Fig.4B, 4D). The S_TKc domain deletion strain (sch9ΔS_TKc), S_TK_X domain deletion strain (sch9ΔS_TK_X), and sch9K340A mutant displayed a similar phenotype to the Δsch9 strain (Fig. 4B, 4C, 4D). These findings revealed that Sch9 plays a significant role in the biological synthesis of aflatoxin and in responding to various stresses. This role is primarily mediated through its kinase domain, kinase-C domain, and ATP binding site k340. We also found that the phosphorylation level of Hog1 increased significantly in the Δsch9 strain under osmotic stress (Fig. 4E). Based on the above results, we hypothesized that Sch9 was involved in the osmotic stress response by modulating the HOG pathway.

Revised version: In S. cerevisiae, the protein kinase Sch9 is a major downstream effector of TORC1, and this kinase is known to have crucial functions in stress resistance, longevity, and nutrient sensing [44]. In addition, the TORC1-Sch9 pathway acts as a crucial mediator of chronological lifespan [44]. In filamentous fungi, an intricate network of interactions exists between the TOR and HOG pathways. In F. graminearum [38], FgSch9 and FgHog1 mutants displayed heightened susceptibility to osmotic and oxidative stresses. The phosphorylation level of FgHog1 increased significantly in FgSch9 mutant. Additionally, the affinity capture-MS assay showed that the Tor kinase was among the proteins co-purifying with FgSch9. These results suggest that FgSch9 serves as an intermediary for the TOR and HOG signaling pathways [38]. Furthermore, FgSch9 plays a crucial role in the regulation of various biological processes such as hyphal differentiation, asexual development, virulence, and DON biosynthesis[38]. In A. fumigatus, the SchA mutant was sensitive to rapamycin, high concentrations of calcium and hyperosmotic stress, and SchA was involved in iron metabolism. The SchA mutant exhibited enhanced phosphorylation of SakA in response to osmotic stress [29]. In this study, we discovered that Sch9 is involved in regulating aflatoxin biosynthesis (Fig. 3C), and responses to osmotic stress and rapamycin stress (Fig. 3B, 3D). The S_TKc domain deletion strain (sch9ΔS_TKc), S_TK_X domain deletion strain (sch9ΔS_TK_X), and sch9K340A mutant displayed a similar phenotype to the Δsch9 strain (Fig. 3B, 3C, 3D). These findings revealed that Sch9 plays a significant role in the biological synthesis of aflatoxin and in responding to various stresses. This role is primarily mediated through its kinase domain, kinase-C domain, and ATP binding site K340. We also found that the phosphorylation level of Hog1 increased significantly in the Δsch9 strain under osmotic stress (Fig. 3E). Based on the above results, we hypothesized that Sch9 was involved in the osmotic stress response by modulating the HOG pathway.

Earlier version: In S. pombe, the protein Tip41, which interacts with Tap42, plays a significant role in the cellular responses to nitrogen sources by regulating type 2A phosphatases [60]. In S. cerevisiae, Tap42-phosphatase complexes associate with TORC1, whereas Tip41 functions to attenuate TORC1 signaling by stimulating the PP2A phosphatase [48]. Unlike its yeast counterpart Tip41, TIPRL, a mammalian Tip41-like protein, exerts a positive influence on mTORC1 signaling by interacting with PP2Ac [61]. While the interaction between Tip41 and Tap42 was not observed in the filamentous fungi. The yeast two-hybrid (Y2H) assay showed that FgTip41 interacts with the phosphatase FgPpg1 rather than FgTap42, while the FgTap42 bind with phosphatases FgPp2A, implying that FgTap42, FgPpg1, and FgTip41 form a heterotrimer in F. graminearum [23]. In the rice blast fungus M. oryzae, MoTip41 does not interact with the MoTap42. However, MoTip41 has the ability to bind with the phosphatase MoPpe1, thereby facilitating crosstalk between the TOR pathway and the cell wall integrity pathway [40]. In this study, we found that the TapA and TipA played a crucial role in sclerotia production and cell wall stress in A. flavus (Fig. 5B, S8). The tapA single-copy deletion mutant and the tipA deletion mutant exhibited decreased sensitivity to rapamycin (Fig. 5A). However, the interaction between TapA and TipA is not yet clear in A. flavus. We speculated that the interrelations and functions of TapA and TipA in A. flavus are comparable to their ortholog in other filamentous fungi.

Revised version: Tap42, a protein that associates with phosphatase 2A, serves as the direct target of the Tor kinase in yeast. Tap42 plays a crucial role in various Tor functions, particularly in the regulation of transcriptional processes [58]. The genome of A. flavus contains a homolog of tap42 named tapA (AFLA_092770), which is predicted to encode a 355 amino-acid protein (sharing 31.9% identity with S. cerevisiae Tap42). To further determine the function of the TapA regulator in the TOR signaling pathway of A. flavus, we conducted targeted gene deletion experiments on tapA. We failed to obtain a null mutant, we hypothesised that tapA knockout was lethal in A. flavus. In S. pombe, the protein Tip41, which interacts with Tap42, plays a significant role in the cellular responses to nitrogen sources by regulating type 2A phosphatases [59]. In S. cerevisiae, Tap42-phosphatase complexes associate with TORC1, whereas Tip41 functions to attenuate TORC1 signaling by stimulating the PP2A phosphatase [45]. Unlike its yeast counterpart Tip41, TIPRL, a mammalian Tip41-like protein, exerts a positive influence on mTORC1 signaling by interacting with PP2Ac [60]. While the interaction between Tip41 and Tap42 was not observed in the filamentous fungi. The yeast two-hybrid (Y2H) assay showed that FgTip41 interacts with the phosphatase FgPpg1 rather than FgTap42, while the FgTap42 bind with phosphatases FgPp2A, implying that FgTap42, FgPpg1, and FgTip41 form a heterotrimer in F. graminearum [23]. In the rice blast fungus M. oryzae, MoTip41 does not interact with MoTap42. However, MoTip41 can bind with the phosphatase MoPpe1, thereby facilitating crosstalk between the TOR pathway and the cell wall integrity pathway [40]. In this study, we found that TipA played a crucial role in sclerotia production and cell wall stress in A. flavus (Fig. 4B, 4H). However, the interaction between TapA and TipA is not yet clear in A. flavus.

Earlier version: In S. cerevisiae, Tap42 can bind and regulate various PP2A phosphatases, including Pph3, Pph21, Pph22, Sit4, and Ppg1. This interaction plays a crucial role in the TOR signaling pathway, particularly in response to rapamycin treatment or nutrient deprivation [48]. The phosphatase Sit4 plays a crucial role in regulating cell growth, morphogenesis, and virulence in C. albicans [62]. FgSit4 and FgPpg1 are associated with various cellular processes, including mycelial growth, conidiation, DON biosynthesis, and virulence in F. graminearum [23]. In addition, FgSit4 and FgPpg1 are involved in the regulation of various signaling pathways, such as the CWI pathway and lipid droplets biogenesis regulation pathway [23]. In this study, we found that the phosphatase SitA and Ppg1 play critical role in vegetative growth, conidiation, sclerotia formation, and aflatoxin biosynthesis in A. flavus (Fig.6A, 6B, 6C). Additionally, the phosphatase SitA and Ppg1 also affect pathogenicity and LD biogenesis (Fig. 7A, 9C, 9D). In addition, TapA/TipA and SitA/Ppg1 exhibits similar functions in sclerotia production and cell wall stress response. Based on these observations, we hypothesized that TapA-phosphatase complexes play a collaborative role in the regulation of the sclerotia formation and cell wall stress response in A. flavus.

Revised version: In S. cerevisiae, Tap42 can bind and regulate various PP2A phosphatases, including Pph3, Pph21, Pph22, Sit4, and Ppg1. This interaction plays a crucial role in the TOR signaling pathway, particularly in response to rapamycin treatment or nutrient deprivation [45]. The phosphatase Sit4 plays a crucial role in regulating cell growth, morphogenesis, and virulence in C. albicans [61]. FgSit4 and FgPpg1 are associated with various cellular processes, including mycelial growth, conidiation, DON biosynthesis, and virulence in F. graminearum [23]. In addition, FgSit4 and FgPpg1 are involved in the regulation of various signaling pathways, such as the CWI pathway and lipid droplets biogenesis regulation pathway [23]. In this study, we found that the phosphatase SitA and Ppg1 play critical role in vegetative growth, conidiation, sclerotia formation, and aflatoxin biosynthesis in A. flavus (Fig. 5A, 5B, 5C). Additionally, the phosphatase SitA and Ppg1 also affect pathogenicity and LD biogenesis (Fig. 6A, 8C, 8D). In addition, SitA, Ppg1, and TipA exhibits similar functions in sclerotia production and cell wall stress response. Based on these observations, we hypothesized that TapA-phosphatase complexes play a collaborative role in the regulation of the sclerotia formation and cell wall stress response in A. flavus.

Earlier version: After the deletion of the sitA and ppg1 genes, the expression level of genes related to cell wall synthesis has decreased significantly, resulting in damage to the cell wall and showing a higher sensitivity to cell wall integrity stress (Fig. 8C). Western blotting results showed that the phosphorylation level of MAPK kinase Slt2 has increased significantly in the ΔsitA and Δppg1 strains under cell wall stress, indicating that SitA and Ppg1 exert a negative regulatory effect on the CWI pathway (Fig. 8D). In M. oryzae, MoPpe1 and MoSap1 serve as an adaptor complex that connects the CWI and TOR signaling pathways. Activation of the TOR pathway results in the inhibition of the CWI pathway [39]. In addition, the phosphatase MoPpe1 interacts with the phosphatase-associated protein MoTip41, facilitating the exchange of information between the TOR and CWI signaling pathways [39]. These results suggest that the phosphatase Ppe1 engages in interactions with various proteins to effectively coordinate the TOR and CWI signaling pathways, thereby regulating the growth and pathogenicity of the rice blast fungus M. oryzae [39]. Taken together, we speculated that the phosphatase SitA and Ppg1 might exhibit differential regulation in the CWI and TOR pathways across species in order to fulfill their respective functions.

Revised version: After the deletion of the sitA and ppg1 genes, the expression level of genes related to cell wall synthesis has decreased significantly, resulting in damage to the cell wall and showing a higher sensitivity to cell wall integrity stress (Fig. 7C). Western blotting results showed that the phosphorylation level of MAPK kinase Slt2 has increased significantly in the ΔsitA and Δppg1 strains under cell wall stress, indicating that SitA and Ppg1 exert a negative regulatory effect on the CWI pathway (Fig. 7D). In M. oryzae, MoPpe1 and MoSap1 serve as an adaptor complex that connects the CWI and TOR signaling pathways. Activation of the TOR pathway results in the inhibition of the CWI pathway [39]. In addition, the phosphatase MoPpe1 interacts with the phosphatase-associated protein MoTip41, facilitating the exchange of information between the TOR and CWI signaling pathways [39]. These results suggest that the phosphatase Ppe1 engages in interactions with various proteins to effectively coordinate the TOR and CWI signaling pathways, thereby regulating the growth and pathogenicity of the rice blast fungus M. oryzae [39]. Taken together, we speculated that the phosphatase SitA and Ppg1 might exhibit differential regulation in the CWI and TOR pathways across species in order to fulfill their respective functions.

Earlier version: In F. graminearum, the phosphatase complex FgNem1/Spo7-FgPah1 cascade plays significant roles in fungal development, DON production, and lipid droplet (LD) biogenesis [24]. In this study, our results indicated that Nem1 and Spo7 are important for vegetative growth and secondary metabolism of A. flavus (Fig. 10A, 10B). We have also discovered that LD biogenesis is facilitated by the Nem1/Spo7 and SitA/Ppg1 pathways (Fig. 9C, 9D, 10E). These observations suggest that the PP2A phosphatases play a central role as core phosphatases in the LD biogenesis signaling network. However, it is likely that there are other kinases involved in the regulation of LD biogenesis. To gain a comprehensive understanding of the precise molecular mechanism, further research should focus on investigating the interplay between phosphatases and other kinases.

Revised version: In F. graminearum, the phosphatase complex FgNem1/Spo7-FgPah1 cascade plays significant roles in fungal development, DON production, and lipid droplet (LD) biogenesis [24]. In this study, our results indicated that Nem1 and Spo7 are important for vegetative growth and secondary metabolism of A. flavus (Fig. 9A, 9B). We have also discovered that LD biogenesis is facilitated by the Nem1/Spo7 and SitA/Ppg1 pathways (Fig. 8C, 8D, 9E). These observations suggest that the PP2A phosphatases play a central role as core phosphatases in the LD biogenesis signaling network. However, there are likely other kinases involved in the regulation of LD biogenesis. To gain a comprehensive understanding of the precise molecular mechanism, further research should focus on investigating the interplay between phosphatases and other kinases.

In conclusion, we identified the complex regulatory network that provides insight into the interplay between the TOR, HOG, and CWI signaling pathways, and how they collectively regulate the development, aflatoxin biosynthesis, and pathogenicity of A. flavus. By elucidating the functions of genes associated with the TOR signaling pathway, our research could establish a solid foundation for comprehending the molecular mechanism underlying the aflatoxin biosynthesis of A. flavus. This, in turn, can provide novel insights into the management of aflatoxin contamination and the mitigation of A. flavus infection.

Materials and methods

Strains and cultural conditions

Earlier version: The A. flavus strains in this study were shown in Table S3. All strains were cultivated on three different types of agar media: potato dextrose agar (PDA), yeast extract-sucrose agar (YES), and yeast extract-glucose agar (YGT). The growth and conidiation assays were conducted at 37℃. Additionally, aflatoxin production was assessed using YES liquid medium at 29℃.

Revised version: The A. flavus strains in this study were shown in Table S1. All strains were cultivated on three different types of agar media: potato dextrose agar (PDA), yeast extract-sucrose agar (YES), and yeast extract-glucose agar (YGT). The growth and conidiation assays were conducted at 37℃. Additionally, aflatoxin production was assessed using YES liquid medium at 29℃.

Constructions of gene deletion and complemented mutants

Earlier version: The construction of gene deletion and complemented strains was carried out using the SOE-PCR strategy, following previously established protocols [63]. The primers utilized for the amplification of the flanking sequences of each gene can be found in Table S3. The putative gene deletion mutants were validated through PCR assays using the appropriate primers (Fig. S2), and their confirmation was subsequently obtained through sequencing analysis.

Revised version: The construction of gene deletion and complemented strains was carried out using the SOE-PCR strategy, following previously established protocols [62]. The primers utilized for the amplification of the flanking sequences of each gene can be found in Table S2. The putative gene deletion mutants were validated through PCR assays using the appropriate primers (Fig. S2), and their confirmation was subsequently obtained through sequencing analysis.

Removed in the revised version: Detection of tor and tapA genes copy number in strains

Removed in the revised version: Absolute quantification PCR has been used as a common method for molecular detection of gene copy numbers. According to previous reports, the sumo gene is a single-copy gene in A. flavus[64]. Following previously established protocols[65], the primers were used to design the amplification of the sumo reference gene and tor and tapA gene. Using sumo gene as reference, the tor and tapA gene copy number was calculated by standard curve.

Growth, conidiation, and sclerotia analysis

Earlier version: To evaluate mycelial growth and conidia formation, 1 μL of 106 spores/mL conidia of all strains was point inoculated onto the center of PDA medium, and then cultured at 37℃ for 5 days in the dark. The conidia were washed by 0.07% Tween-20 and counted using a hemocytometer and microscope [43]. For sclerotia formation analysis, each strain was inoculated and cultivated on CM agar medium at 37℃ in the dark for 7 days, and then 75% ethanol was used to wash away mycelia and conidia on the surface of the medium [64]. These assays were repeated at least three times.

Revised version: To evaluate mycelial growth and conidia formation, 1 μL of 106 spores/mL conidia of all strains was point inoculated onto the center of the PDA medium, and then cultured at 37℃ for 5 days in the dark. The conidia were washed with 0.07% Tween-20 and counted using a hemocytometer and microscope [43]. For sclerotia formation analysis, each strain was inoculated and cultivated on CM agar medium at 37℃ in the dark for 7 days, and then 75% ethanol was used to wash away mycelia and conidia on the surface of the medium [63]. These assays were repeated at least three times.

Extraction and quantitative analysis of aflatoxin

Earlier version: To extract and quantify aflatoxins (AFs) production, 1 μL of 106 spores/mL conidia of all strains was point inoculated onto liquid YES medium and cultures were then incubated at 29℃ in the dark for 6 days. AFs were extracted from the media and detected using thin layer chromatography (TLC) [65]. In addition, the standard AFB1 was employed for visual comparison, and AFs were quantitatively analyzed using the Image J tool. The assays were repeated at least three times.

Revised version: To extract and quantify aflatoxins (AFs) production, 1 μL of 106 spores/mL conidia of all strains was point inoculated onto liquid YES medium, and cultures were then incubated at 29℃ in the dark for 6 days. AFs were extracted from the media and detected using thin layer chromatography (TLC) [63]. In addition, the standard AFB1 was employed for visual comparison, and AFs were quantitatively analyzed using the Image J tool. The assays were repeated at least three times.

Stress assay

To determine the roles of the gene in A. flavus responsing to various environmental stresses, 1 μL of 106 spores/mL conidia of all strains were point inoculated onto these PDA solid plates supplemented with various agents. These agents included 1M NaCl, 1M KCl, and 1M CaCl2 for inducing osmotic stress, as well as 100 µg/mL SDS, 200 µg/mL Calcofluor white (CFW), and 200 µg/mL Congo red (CR) for inducing cell wall stress. After 5 days of incubation at 37℃ in the dark, the relative rates of inhibition were determined [43]. The assays were repeated at least three times.

Seeds infection assays

To measure the pathogenicity of A. flavus, all peanut and maize seeds were washed in 0.05% sodium hypochlorite, 75% ethanol, and sterile water, respectively. Subsequently, all peanut and maize seeds were immersed in sterile water containing 106 spores/mL of all strains at 29℃ in the dark for 60 min. After the conidia adhesion, the seeds were cultured at 29℃ for 6 days in the dark [64].

Quantitative RT-PCR

Earlier version: For the real-time quantitative PCR (qPCR) assays, the A. flavus strains were cultivated in PDA medium at 29℃ for 3 days in the dark. Total RNA was extracted from the mycelia of all strains, and cDNA was synthesized with the First Strand cDNA Synthesis Kit, and then the cDNA was used as the template for qRT-PCR analysis with SYBR Green qPCR mix. The relative transcript level of each gene was calculated by the 2-ΔΔCTmethod [66].

Revised version: For the real-time quantitative PCR (qPCR) assays, the A. flavus strains were cultivated in the PDA medium at 29℃ for 3 days in the dark. Total RNA was extracted from the mycelia of all strains, and cDNA was synthesized with the First Strand cDNA Synthesis Kit, and then the cDNA was used as the template for qRT-PCR analysis with SYBR Green qPCR mix. The relative transcript level of each gene was calculated by the 2-ΔΔCTmethod [65].

Western blot analysis

To extract the total protein, the mycelia of each strain were cultivated in YES liquid medium for 48 h, then the sample was grinded under liquid nitrogen to a fine powder and transferred to a 1.5 mL microcentrifuge tube added with RIPA lysis buffer and 1 mM PMSF protease inhibitor. Proteins were separated by 12% SDS-PAGE and transferred onto a polyvinylidene fluoride (PVDF) membrane. Phosphorylation levels of Hog1 and Slt2 were detected by Phospho-p38 MAPK antibody (Cell Signaling Technology, Boston, MA, USA) and Phospho-p44/42 MAPK (Erk 1/2) antibody (Cell signaling Technology, MA, USA), respectively. The experiment was conducted three times independently.

Statistical analysis

Earlier version: All data were measured with means ± standard deviation (SD) of three biological replicates. Graph Pad Prism 7 software was used for conducting statistical and significance analysis, and p-values < 0.05 was represented statistical significance. Student’s t-test was used to compare the two means for differences, whereas Tukey’s multiple comparisons were used for testing multiple comparisons.

Revised version: All data were measured with means ± standard deviation (SD) of three biological replicates. Graph Pad Prism 7 software was used for conducting statistical and significance analysis, and p-values < 0.05 represented statistical significance. Student’s t-test was used to compare the two means for differences, whereas Tukey’s multiple comparisons were used for testing multiple comparisons.

Supporting information